DePARylation is critical for S phase progression and cell survival

Author Response

The following is the authors’ response to the previous reviews.

Reviewer #2 (Recommendations For The Authors):

I would like to thank the authors for their comments. However, my request for additional experiments to consolidate this manuscript and text changes have not been addressed (point 1 and point 2), which I believe are essential for completion of this manuscript.

The reviewer raised the question about the relevant substrates of PARG in S-phase cells (point 1). As we explained in our previous response, the most important substrate of PARG is PARP1, since we observed increased chromatin-associated PARP1 and PARylated PARP1 in cells with PARG depletion. Moreover, PARP1 or PARP1/2 depletion rescued cell lethality caused by PARG depletion. These data strongly suggest that PARP1 is the major substrate of PARG in S phase cells. Of course, PARG may have additional substrates. In the future, we will perform proteomics experiments as suggested by this reviewer to identify additional PARG substrates, which may reveal new roles of PARG in S phase progression.

The reviewer also suggested us to re-organize our manuscript (point 2). However, we prefer to keep the manuscript as it is, since this is how the project evolved. The other reason we would like to share with the readers is the challenge to validate KO cells. This is an important lesson we learned from this study. We hope that this will raise the awareness of hypomorphic mutant cells we often use to draw conclusions about gene functions and/or genetic interactions. We understand that the current flow of our manuscript may bring some confusion. To avoid it, we included additional explanations at the beginning of this manuscript to draw attention to the readers that our initial KO cells may not be complete PARG KO cells, i.e. they may have residual PARG activity. We also included additional discussion of this important point in the Discussion section.

Moreover, WB analysis of PARG KO clones is inconclusive, as the additional prominent band at 50 kDa could be a degradation product. The authors should check PARG levels are localization by IF, which allows detection of intact proteins and their cellular localizations, since the shorter isoform should be localized in the cytosol. WB with PARG isoforms is missing important information regarding Mw of the PARG constructs and Mw labels of western blots, which makes is difficult to evaluate this data and compare to KO. Ideally, KO and PARG isoform samples should be all on one gel for proper comparison with different antibodies.

We appreciate the concerns raised by this reviewer. We agree that the additional prominent band at 50kDa could be a degradation product. As we explained in our previous response, despite using several PARG antibodies, we could not draw a clear conclusion which functional isoforms or truncated forms were expressed in our PARG KO cells.

Immunostaining experiments may not be more conclusive, since IF experiments rely on the same antibodies for recognizing endogenous PARG. Additionally, even a protein mainly localizes in the cytosol, we cannot exclude the possibility that a small fraction of this protein may localize in nuclei and have nuclear functions.

Instead, as we presented in our manuscript, we used a biochemical assay to measure PARG activity in cell lysate and showed that our initial PARG KO cells still have residual PARG activity. However, we could not detect any PARG activity in our complete/conditional PARG KO cells (cKO cells; these cells can only survive in the presence of PARP inhibitor). These data strongly suggest that PARG is essential for cell survival.

Changes from the earlier version to the revised version

Abstract

Abstract

Poly(ADP-ribose)ylation or PARylation by PAR polymerase 1 (PARP1) and dePARylation by poly(ADP-ribose) glycohydrolase (PARG) are equally important for the dynamic regulation of DNA damage response. PARG, the most active dePARylation enzyme, is recruited to sites of DNA damage via pADPr-dependent and PCNA-dependent mechanisms. Targeting dePARylation is considered an alternative strategy to overcome PARP inhibitor resistance. However, precisely how dePARylation functions in normal unperturbed cells remains elusive. To address this challenge, we conducted multiple CRISPR screens and revealed that dePARylation of S phase pADPr by PARG is essential for cell viability. Loss of dePARylation activity initially induced S phase-specific pADPr signaling, which resulted from unligated Okazaki fragments and eventually led to uncontrolled pADPr accumulation and PARP1/2-dependent cytotoxicity. Moreover, we demonstrated that proteins involved in Okazaki fragment ligation and/or base excision repair regulate pADPr signaling and cell death induced by PARG inhibition. In addition, we determined that PARG expression is critical for cellular sensitivity to PARG inhibition. Additionally, we revealed that PARG is essential for cell survival by suppressing pADPr. Collectively, our data not only identify an essential role for PARG in normal proliferating cells but also provide a potential biomarker for the further development of PARG inhibitors in cancer therapy.

Introduction

Poly(ADP-ribosyl)ation or PARylation is a conserved post-translational modification that is important for many cellular processes, including DNA damage repair1, 2. PARylation is mainly driven by poly(ADP-ribose) polymerase 1 (PARP1), and to a lesser extent by PARP2, using NAD+ as an ADP-ribose donor and generating nicotinamide (NAM) as a byproduct in human cells. Extensive studies have already established the key roles of PARP1/2 in DNA damage response (DDR)3, 4. As a DNA damage sensor, PARP1 rapidly recognizes and binds to DNA damage sites, which dramatically activates its own enzymatic activity and thereby modifies itself and other proteins with pADPr and mono-(ADP-ribosylation). These modifications lead to the recruitment of proteins that are involved in DDR and DNA repair. The binding to DNA by activated PARP1 acts as a “hit and run” mechanism5, which restricts PARylation at a specific point and facilitates the subsequent engagement of other repair proteins. Thus, PARylation by PARP1, when responding to DNA damage, is a rapid and transient process. In addition, PARP1 can be activated by unligated Okazaki fragments to induce endogenous S phase pADPr, which then recruits XRCC1 and LIG3 to facilitate the ligation of these Okazaki fragments6.

PARylation is a reversible post-translational modification. The removal of pADPr is mainly carried out by poly(ADP-ribose) glycohydrolase (PARG)7. However, PARG cannot remove the terminal ADP-ribose, whose removal requires additional hydrolases, including terminal ADP-ribose glycohydrolase (TARG1), ADP-ribose-acceptor hydrolases ARH1/3, and possibly other macrodomain-containing proteins, such as hMacroD1/D28. Later studies showed that PARG and ARH3 (ADPRHL2/ADPRS) are the primary dePARylation enzymes in vertebrates, although ARH3 has much lower activity against pADPr than PARG and mainly functions as a serine-directed mono-ADP-ribosylhydrolase9–11. As an endo-glycohydrolase and exo-glycohydrolase, PARG specifically hydrolyzes the glycosidic bonds. PARG may function as an oncogene, as the high level of PARG promotes cell transformation and invasion and is associated with poor overall survival12. Moreover, complete loss of PARG leads to embryonic lethality in mice13. Importantly, PARG is recruited to sites of DNA damage by PARP1- or pADPr -dependent and PCNA-dependent mechanisms14 and involved in both DNA double-strand breaks (DSBs) and single-strand breaks (SSBs) repair. Loss of PARG sensitizes cells to DNA damage agents15–17. In addition, PARG is also involved in transactions at replication forks18, 19.

On average, a single cell suffers thousands of DNA damage events every day20; PARylation and dePARylation occur rapidly and collaboratively to facilitate DNA repair in the cell. This rapid PARylation and dePARylation cycle is likely important for homeostasis regulation. Unrestrained PARylation following DNA damage may lead to a significant change in the NAD+ level in cells, which not only inhibits other cellular processes that require NAD+ but also results in cell death due to NAD+ depletion and/or PARP1-dependent cell death called Parthanatos10, 21. Thus, PARG and other dePARylation enzymes are critically important for the recycling of NAD+ and the control of NAD+ homeostasis22–24.

Earlier version: Targeting PARylation and dePARylation is a promising strategy in cancer and other therapies. PARP inhibitors (PARPis) have already been approved for clinical treatment of cancers with homologous recombination (HR) deficiency, while targeting dePARylation is an encouraging alternative strategy to overcome PARPi resistance25–29. Given the critical roles of PARG in the maintenance of the PARylation/dePARylation cycle in cells, PARG inhibitors have been developed recently as potential anti-cancer agents28. These include PDD00017273 (PARGi) (Fig. S1A)30, COH3426, and JA213131. Interestingly, cells with deficiency in HR and other DNA damage signaling/repair pathways showed increased sensitivity to the PARG inhibitors PDD00017273 and COH3425, 26. These PARG inhibitors could also sensitize cells to radiation therapy and chemotherapy23, 26, 31, 32. Thus, there is significant interest in further developing these PARG inhibitors for cancer treatment.

Revised version: Targeting PARylation and dePARylation is a promising strategy in cancer and other therapies. PARP inhibitors (PARPis) have already been approved for clinical treatment of cancers with homologous recombination (HR) deficiency, while targeting dePARylation is an encouraging alternative strategy to overcome PARPi resistance25–29. Given the critical roles of PARG in the maintenance of the PARylation/dePARylation cycle in cells, PARG inhibitors have been developed recently as potential anti-cancer agents28. These include PDD00017273 (PARGi) (Figure 1-figure supplement 1A)30, COH3426, and JA213131. Interestingly, cells with deficiency in HR and other DNA damage signaling/repair pathways showed increased sensitivity to the PARG inhibitors PDD00017273 and COH3425, 26. These PARG inhibitors could also sensitize cells to radiation therapy and chemotherapy23, 26, 31, 32. Thus, there is significant interest in further developing these PARG inhibitors for cancer treatment.

Despite the crucial role of PARG in dePARylation in multiple pathways, detailed knowledge of its mechanism of action remains elusive. Especially, how PARG regulates dePARylation in normal unperturbed cells and the balance between PARylation and dePARylation remain unclear. To address this issue, we performed multiple CRISPR screens and created PARP1/2 DKO, PARG KO, and additional KO cells. We showed that loss of dePARylation activity induced S phase-specific pADPr signaling, which likely originated from unligated Okazaki fragments. Consequently, the failure to remove S phase pADPr led to PARylation- or PARP1/2-dependent cell death. Furthermore, perturbation of Okazaki fragment ligation and/or base excision repair (BER) increased pADPr signaling and promoted cell death induced by PARGi. In addition, we revealed that the PARG level is a potential biomarker for PARGi-based cancer therapy. Moreover, we showed that PARG is essential for cell survival, which can be exploited for cancer treatment.

Added in the revised version: Results

Added in the revised version: PARG depletion leads to drastic sensitivity to PARGi

Added in the revised version: The results of early studies indicated that defective DNA repair pathways and DNA replication stress would result in enhanced sensitivity to PARG inhibitor, PDD00017273 (PARGi)25–27, 33, 34. To further define the key DDR pathways that are important for cellular response to PARGi, we performed CRISPR screening using our homemade DDR sgRNA library, which targets approximately 360 genes involved in various DDR pathways35. Interestingly, our screen results showed that PARG depletion led to significantly increased cellular sensitivity to PARGi (Figure 1A, Supplementary File 1a). In addition, we found that POLB (Figure 1A), which is important for BER, a repair pathway that relies on PARP1 functions and pADPr-dependent DNA damage signaling, also showed synthetic lethality with PARGi. The synthetic lethality between POLB and PARGi was consistent with the findings of a previous report33, indicating the high quality of our screening results.

Added in the revised version: PARG loss sensitizes cells to PARGi.

Added in the revised version: a, Ranking of co-essential genes with PARGi treatment on the basis of a DrugZ analysis of the results of CRISPR/Cas9 screening with a DDR library in HEK293A cells. The NormZ score was used to determine a possible synthetic lethality gene under PARGi treatment. Drug-sensitive genes were marked in red; drug-resistant genes were marked in blue on the basis of the false discovery rate (FDR, 0.05 cut-off). b, HEK293A WT cells and HEK293A PARG KO cells were treated with different doses of PARGi for 72 hours. Cell viability was determined by the CellTiter-Glo assay. c, HEK293A PARG KO cells, re-constituted with either full-length PARG or catalytic domain mutation of PARG, were treated with different doses of PARGi for 72 hours. Cell viability was determined by the CellTiter-Glo assay.

Added in the revised version: To further validate our data, we generated PARG KO cells in the same 293A cells used in DDR sgRNA library screening. Interestingly, PARG KO cells showed extreme sensitivity to PARGi, with nearly thousand-fold sensitivity (the IC50 in 293A cells was 96±24 µM and 210±30 nM in PARG KO cells, on the basis of five independent biological replicates) (Figure 1B). In addition, we used another two structurally similar compounds (PDD00017272 and PDD00017238) to further confirm the extreme sensitivity of PARG KO cells to PARG inhibition (Figure 1-figure supplement 1B), while the inactive but structurally similar compound PDD0031705 did not distinguish wild-type (WT) or PARG KO cells (Figure 1-figure supplement 1C). To further validate that this sensitivity was not restricted to one cell line, we generated HeLa-derived PARG KO cells with a different gRNA. As expected, HeLa-derived PARG KO cells also displayed extreme sensitivity to PARGi (Figure 1-figure supplement 1D). These results suggest that PARG depletion resulted in drastic sensitivity to PARGi. We reconstituted PARG KO cells with WT PARG or a PARG catalytic-inactive mutant E755/756A36 to further investigate the nature of PARGi sensitivity in PARG KO cells. Cells treated with the alkylating agent methyl methanesulfonate (MMS), which is known to activate pADPr signaling, were used to test PARG activity. As shown in Figure 1-figure supplement 1E, PARG KO cells and cells reconstituted with PARG catalytically inactive mutant showed a significant increase in pADPr signaling induced by MMS, while WT cells or cells reconstituted with WT PARG did not. Expectedly, reconstitution with WT PARG reversed the sensitivity of PARG KO cells to PARGi, while reconstitution with PARG catalytic inactive mutant failed to do so (Figure 1C). Together, these data strongly suggest that loss of PARG activity induces dramatic sensitivity to PARGi.

Added in the revised version: The question is why PARG KO cells would be sensitive to PARGi. There are at least two potential explanations for this observation. First, PARGi may target not only PARG but also a second target, which only becomes essential in the absence of PARG. Second, our PARG KO cells, despite being generated with two different sgRNAs and in two independent cell lines validated by Western blotting and DNA sequencing, are not complete PARG KO cells.

Added in the revised version: There are five reviewed or potential PARG isoforms identified in the Uniprot database. The two different sgRNAs used here target all three catalytically active isoforms (isoforms 1, 2 and 3), and sgRNA#2 used in HeLa cells also targets isoforms 4 and 5, but these two isoforms are considered catalytically inactive according to the Uniprot database (Figure 1-figure supplement 1D). Nevertheless, it is possible that sgRNA-mediated genome editing may lead to the creation of new alternatively spliced PARG mRNAs and/or the use of alternative ATG as start codon, which can produce residual but catalytically active forms of PARG. It would be challenging to search for these putative PARG isoforms due to limited PARG antibodies and conceivable low expression of these isoforms. To address this important question, we performed several experiments to test both hypotheses. Please see below for additional studies.

Added in the revised version: PARGi induces PARP1/2-dependent cell death in PARG KO cells

Added in the revised version: As shown above (Figure 1-figure supplement 1E), MMS treatment significantly increased pADPr signaling in PARG KO cells. Interestingly, pADPr signaling further increased modestly in PARG KO cells treated with the combination of PARGi and MMS (Figure 2-figure supplement 1A). To further evaluate the underlying mechanisms of PARGi sensitivity in PARG KO cells, we generated PARP1/2 DKO cells in both WT and PARG KO cells and treated them with PARGi (Figure 2A). PARP1 cleavage was detected in PARG KO cells treated with PARGi, indicating apoptosis. Moreover, pADPr accumulated significantly in PARG KO cells treated with PARGi, but not in other cells (Figure 2-figure supplement 1B). Furthermore, we used flow cytometry to detect pADPr signaling following short treatment with PARGi (Figure 2B). Consistently, moderate pADPr signaling was observed in PARG KO cells, but not in other cells (Figure 2B).

Added in the revised version: PARGi treatment induces NAD+- and PARP-dependent cell death in PARG KO cells.

Added in the revised version: a, HEK293A WT, PARG KO, PARP1/2 DKO, and PARG/PARP1/2 TKO cells were treated with PARGi (1µM) for 72 hours. The total cell lysates were immunoblotted with the indicated antibodies. b, HEK293A WT, PARG KO, PARP1/2 DKO, and PARG/PARP1/2 TKO cells were treated with DMSO or 10 µM PARGi for 4 hours and then fixed and stained with anti-pADPr antibody and propidium iodide (PI). c, HEK293A WT, PARG KO, PARP1/2 DKO clls, and PARG/PARP1/2 TKO cells were treated with different doses of PARGi for 72 hours. Cell viability was determined by the CellTiter-Glo assay. d, Relative NAD+ level in HEK293A WT and PARG KO cells with the indicated treatment for 48 hours. PARGi, 10 µM; PARPi, 10 µM; NAM, 100 µM; NMN, 1 mM; FK866, 10 nM. e, PARG KO cells were treated with PARGi (10 µM) or PARGi and NAM (100 µM) or NMN (1mM) for 48 hours. The total cell lysates were immunoblotted with the indicated antibodies. f, Results of clonogenic assays conducted using HEK293A PARG KO cells treated with PARGi (500 nM) or PARGi and NAM (100 µM) or NMN (1 mM) for 7 days.

Added in the revised version: We reason that sensitivity to PARGi in PARG KO cells may depend on PARP1/2-dependent pADPr accumulation. To test this hypothesis, CellTiter-Glo and a colony formation assay were used to detect the sensitivity of WT, PARG KO, PARP1/2 DKO, and PARG/PARP1/2 TKO cells to PARGi. As expected, the sensitivity of PARG KO to PARGi was completely reversed in TKO cells, and DKO cells showed moderate resistance compared with WT cells (Figure 2C). Indeed, TKO cells completely rescued cell death induced by PARGi in PARG KO cells (Figure 2-figure supplement 1C). In addition, treatment with PARPi olaparib was able to reverse PARGi sensitivity in PARG KO cells (Figure 2-figure supplement 1C). These data together suggest that the cytotoxicity of PARGi in PARG KO cells was due to uncontrolled PARylation, which requires PARP1/2.

Added in the revised version: PARP1/2 catalyzes the pADPr reaction by transferring the ADR-ribose moiety of NAD+ to the acceptor proteins with the release of NAM37. Uncontrolled PARylation may lead to cytotoxicity, at least in part because of NAD+ depletion23, 38, 39. To determine whether cell death induced by PARGi in PARG KO cells is at least partially caused by NAD+ depletion, we measured the relative NAD+ level in WT and PARG KO cells treated with PARGi and used the nicotinamide phosphoribosyltransferase (NAMPT) inhibitor FK866 as a positive control. FK866 is a noncompetitive inhibition of NAMPT, a key enzyme involved in the regulation of NAD+ biosynthesis from its natural precursor NAM 40. As expected, the NAD+ level decreased modestly over the treatment period in PARG KO cells, but this decrease was not as dramatic as those observed in cells treated with FK866 (Figure 2-figure supplement 1D). Moreover, treatment with FK866, which significantly decreased the NAD+ level, did not lead to cell death (Figure 2-figure supplement 1E), suggesting that the decreased NAD+ level is not sufficient or the only reason for the cytotoxicity observed in PARG KO cells treated with PARGi.

Added in the revised version: Furthermore, we treated cells with the NAD+ precursors nicotinamide mononucleotide (NMN) or NAM23, 41. As expected, NMN, NAM, and PARPi rescued the NAD+ decrease in both non-treated and PARGi treated cells (Figure 2D). In addition, NMN and NAM were able to not only reduce pADPr accumulation induced by long-term PARGi treatment (Figure 2E), but also rescue cell lethality caused by PARGi treatment in both HEK293A-derived and HeLa-derived PARG KO cells (Figure 2F, Figure 2-figure supplement 1F). Together, these results suggest that uncontrolled PARylation in the absence of any dePARylation activities would lead to cytotoxicity, probably due to several mechanisms including apoptosis and are not limited to NAD+ depletion/reduction.

Added in the revised version: Unexpected S phase pADPr signaling observed in PARG KO cells treated with PARGi

Added in the revised version: PARP1/2 are activated by DNA damage to generate pADPr signals. Under normal growth conditions, PARP1/2 activities may be low; therefore, it was extremely difficult to detect pADPr. Unexpectedly, pADPr signaling from normal proliferating cells without any exogenous DNA damage was detected in PARG KO cells after PARGi treatment (Figure 2B, Figure 2-figure supplement 1B), which was apparent in S phase cells, as detected by FACS analysis (Figure 2B). To further characterize pADPr and DNA damage signaling induced by PARGi or other DNA damaging agents, we used flow cytometry to detect both pADPr and γH2AX signaling (Figure 3A, Figure 3-figure supplement 1A). Interestingly, pADPr signaling was specifically detected in S phase in PARG KO cells treated with PARGi, while at the same time, DNA damage signaling only showed a slight increase or was not observed, as revealed by anti-γH2AX antibody (Figure 3A). On the other hand, MMS treatment led to increased pADPr signaling throughout the cell cycle (Figure 3A). However, it only led to an S phase-specific increase in DNA damage signaling, i.e. γH2AX. Our interpretation of the MMS results is that MMS treatment would lead to DNA alkylation throughout the cell cycle, which is repaired by the BER pathway. BER creates single-strand nicks as repair intermediates, which are recognized by PARP1 and activate PARP1. Therefore, pADPr signaling can be detected in all cell cycle phases in PARG KO cells under this condition. However, single-strand nicks are not recognized by DNA damage checkpoint pathways. These nicks, when encountering replication forks, will be converted to DSBs and therefore activate DNA damage checkpoint kinases, resulting in increased γH2AX signaling specifically in S phase cells.

Added in the revised version: PARGi treatment induces S phase-specific pADPr signaling in PARG KO cells.

Added in the revised version: a, HEK293A WT and PARG KO cells were treated with DMSO or 10 µM PARGi for 4 hours or 0.01% MMS for 30 min and then fixed and stained with anti-pADPr antibody or anti-γH2A.X antibody and PI. b, HEK293A WT and HEK293A PARG KO cells were mock-treated or pre- treated with 2 µM emetine for 90 min and then treated with PARGi for an additional 4 hours. Cells were fixed and stained with anti-pADPr antibody and PI. c, HEK293A PARG KO cells were synchronized with double thymidine block (DTB). Cells remained with DTB or were released from DTB, treated with 10 µM PARGi for 4 hours, and then fixed and stained with anti-pADPr antibody and PI. d, Representative images and results (left) of clonogenic assays conducted using control cells and DTB synchronized or released HEK293A PARG KO cells treated with the indicated doses of PARGi for 7 days, and quantification of crystal violet staining assay (right). e, Results of clonogenic assays were conducted in PARG KO cells with indicated treatment for 7 days (PARGi, 1 µM; adarotene, 200 nM; CD437, 800 nM).

Added in the revised version: To further explore the relationship between pADPr and γH2AX signaling, we conducted a similar flow cytometry analysis in parental WT and PARG KO cells following treatments with different DNA damaging agents. Again, PARGi treatment in PARG KO cells led to an S phase-specific increase in pADPr signaling; this unique pattern was not detected in response to any of the DNA damaging agents used in this study (Figure 3-figure supplement 1A). Our current working hypothesis is that PARG KO cells treated with PARGi may completely block any dePARylation reactions. The dramatic S phase-specific pADPr signaling detected under this condition indicates that PARP1 is specifically activated in S phase cells.

Added in the revised version: During DNA replication, maturation of Okazaki fragments involves single-strand nicks on DNA, which is normally ligated by LIG1. These unligated Okazaki fragment intermediates could activate PARP1 to generate endogenous S phase pADPr6. It is possible that in such situations, the maturation of Okazaki fragments may require PARP1-dependent recruitment of XRCC1/LIG36, which agrees with the results of early studies that indicate that both LIG1 and LIG3 are required for Okazaki fragment maturation42. Of course, PARP1 may have other functions in the S phase, including DNA repair43.

Added in the revised version: The S phase pADPr signaling observed in PARG KO cells after PARGi treatment was observed under conditions in which no or very limited DNA damage or replication stress– induced signaling was detected by anti-γH2AX antibody (Figure 3A) and S phase pADPr resulting from unligated Okazaki fragments were observed under similar conditions6; thus, we reasoned that unligated Okazaki fragments in normal proliferating cells may render this unexpected S phase pADPr signaling in PARG KO cells after PARGi treatment. Indeed, pre-treatment with emetine, which diminishes Okazaki fragments by its anti-protein synthetic activity6, 44, 45, greatly inhibited S phase pADPr signaling in PARG KO cells (Figure 3B).

Added in the revised version: Furthermore, mild inhibition of POLA1 with two inhibitors (i.e. adarotene and CD437) also abolished the PARGi induced S phase pADPr signaling in PARG KO cells (Figure 3-figure supplement 1B). To further support S phase-specific pADPr signaling resulting from unligated Okazaki fragments, we performed double thymidine block (DTB) and release experiments. Indeed, release into S phase is critical for the pADPr signaling observed in PARG KO cells treated with PARGi (Figure 3C). Moreover, the cytotoxicity of PARGi also requires S phase progression, since both control proliferating cells and double thymidine blocked and released cells died following PARGi treatment, while cells arrested by DTB failed to do so (Figure 3D). Additionally, mild inhibition of POLA1 partially rescued the cytotoxicity of PARGi in PARG KO cells (Figure 3E).

Added in the revised version: PARG KO cells show prolonged PARP1 chromatin binding

Added in the revised version: Following DNA damage, PARP1 detects and is activated by both SSBs and DSBs to initiate subsequent DNA repair46, 47. After that, PARP1 dissociates from DNA. If this does not occur, the persistent PARP1-DNA complexes could trap PARP1 on chromatin, which is one of the key characteristics of PARPi-mediated cytotoxicity48. Given that the sensitivity of PARG KO cells to PARGi resulted from PARP1/2-dependent pADPr accumulation, which indicates an association and activation of PARP1/2 by DNA, we used a previously described trapping assay to measure the levels of chromatin-bound PARP148, 49. Unexpectedly, the levels of chromatin-bound PARP1/2 increased in PARG KO cells, regardless of treatment, while the soluble fraction of PARP1 decreased compared with that in WT cells (Figure 4-figure supplement 1A).

Added in the revised version: To further explore the chromatin-bound PARP1, we performed similar soluble and chromatin fractionation and enrichment of PARylated proteins by Af1521 beads. Consistently, chromatin-bound PARP1 increased in PARG KO cells, regardless of whether these cells were treated with PARGi (Figure 4A). Moreover, PARylated PARP1 also increased in chromatin fraction. Interestingly, following PARGi treatment, PARG KO cells showed more pADPr levels than did WT cells in chromatin fraction, while similar pADPr levels were detected in soluble fractions (Figure 4A). Indeed, a recent study by Gogola and colleagues suggest that PARG depletion does not enhance PARP1 dissociation from DNA but prevents excessive PARP1 binding, on the basis of the results of a similar trapping assay and a laser-induced DNA damage assay for measuring PARP1 association50. Taken together, these results indicate that prolonged PARP1 chromatin binding or trapping occurs in PARG KO cells, which may contribute to the sensitivity of these cells to PARGi.

Added in the revised version: Prolonged PARGi treatment induces pADPr throughout the cell cycle and DDR in PARG KO cells.

Added in the revised version: a, Immunoblots of chromatin-bound PARP1 and PARylated proteins in HEK293A WT and PARG KO cells treated with PARGi (10 µM) for 4 hours. PARylated proteins were enriched by Af1521 beads. b, Sensitivity of HEK 293A PARG/PARP1/2 TKO and PARP1 reconstitution cells to PARGi. Cells were treated with different doses of PARGi for 72 hours, and cell viability was determined by the CellTiter-Glo assay. c, Prolonged PARGi treatment induces pADPr throughout the cell cycle in PARG KO cells. HEK293A WT and PARG KO cells were treated with PARGi (10 µM) for the indicated time and then fixed and stained with anti-pADPr antibody and violet. d, Immunoblotting of γH2A.X signals and other indicated proteins and modifications induced by prolonged PARGi+/-MMS treatment.

Added in the revised version: To further support the contribution of chromatin-bound PARP1 to PARGi sensitivity, we introduced either WT PARP1 or a trapping-deficient PARP1del.p.119K120S construction into TKO cells (Figure 4-figure supplement 1B). PARP1del.p.119K120S mutant is a trapping-deficient mutant, which leads to PARPi resistance51, 52. Notably, expression of WT PARP1 led to dramatic PARGi sensitivity in TKO cells, while PARP1del.p.119K120S mutant only showed a slight increase in PARGi sensitivity (Figure 4B). Together, these data suggest that increased chromatin-bound and trapped PARP1 contributes, at least in part, to the sensitivity of these cells to PARGi.

Added in the revised version: To investigate the consequence of S phase pADPr and prolonged PARP1 chromatin binding in PARG KO cells, we again used flow cytometry to detect pADPr signaling in a time course following PARGi treatment. Consistent with the aforementioned data, S phase pADPr was detected following short PARGi treatment (i.e. 4 hours), while pADPr signaling was detected throughout the cell cycle following prolonged PARGi treatment, i.e. 24 or 48 hours (Figure 4C). Furthermore, SSBs were not detected following short PARGi treatment, as measured by an alkaline comet assay (Figure 4-figure supplement 1C), which may due to the sensitivity of this assay. However, prolonged PARGi treatment in PARG KO cells led to dramatic increase of SSBs detected by this assay (Figure 4-figure supplement 1C). Similarly, we observed S phase-specific pADPr signaling with short PARGi treatment and pADPr signaling throughout the cell cycle with prolonged PARGi treatment in HeLa PARG KO cells, although significant S-phase pADPr was detected in control WT HeLa cells (Figure 4-figure supplement 1D), which was not observed in control WT 293A cells (Figure 4C). The difference in pADPr signaling between HeLa and 293A cells may due to a low level of PARG and/or high activity of PARP1 in HeLa cells, in which high–molecular weight smears were detected following PARGi and MMS treatment (Figure 4-figure supplement 1E). Together, these data indicate that PARG regulates S-phase pADPr in normal proliferating cells.

Added in the revised version: To further confirm pADPr signaling observed by flow cytometry analysis, we detected pADPr and γH2AX by Western blotting in HEK293A and HeLa cells and their corresponding PARG KO cells, while additional MMS treatment was included as a positive control (Figure 4D, Figure 4-figure supplement 1F). Expectedly, we only observed increased γH2AX signaling in PARG KO cells with prolonged PARGi treatment (48 hours of treatment in 293A PARG KO cells and 24 and 48 hours of treatment in HeLa PARG KO cells) or those treated with additional MMS, indicating that pADPr accumulation leads to DNA damage throughout the cell cycle. The pADPr signaling in PARG KO cells pre-treated with PARGi was further enhanced in MMS-treated cells (Figure 4D, Figure 4-figure supplement 1F), likely due to increased SSBs; nevertheless, PARGi-mediated pADPr accumulation in PARG KO cells under condition without any exogenous DNA damage was sufficient to cause cytotoxicity. This situation is different from the hyperactivity and progressive inactivation of PARP1 in XRCC1-deficient cells after MMS treatment38. Taken together, we speculate that inhibition of dePARylation in PARG KO cells treated with PARGi lead to pADPr accumulation and cell lethality.

Added in the revised version: CRISPR screens reveal genes responsible for regulating pADPr signaling and/or cell lethality in WT and PARG KO cells

Added in the revised version: PARG KO and PARGi appear to have a synergic effect on cell viability, which depends on PARP1/2-dependent pADPr accumulation (Figure 1, Figure 2). Further, we showed that unligated Okazaki fragments were the likely source of SSBs, which leads to S phase pADPr signaling (Figure 3). Moreover, we found that uncontrolled S phase pADPr accumulation eventually leads to DNA damage, as revealed by anti-γH2AX antibody and alkaline comet assay (Figure 4). These data prompted us to further explore potential genes that may be responsible for controlling cell viability and pADPr signaling in control WT and PARG KO cells, with or without PARGi treatment. As we speculated above, there may be a potential second gene targeted by PARGi, which only becomes essential in the absence of PARG. We thus performed the outlined whole-genome CRISPR screens to address these questions (Figure 5A), similar to the studies published by us and others53, 54.

Added in the revised version: CRISPR screening identifies regulators of pADPr and cell viability.

Added in the revised version: a, Workflow of whole-genome CRISPR screens. For FACS-based CRISPR screening, 5 days after puromycin selection, cells were treated with PARGi (10uM) for 4 hours and then stained with anti-pADPr antibody and sorted with flow cytometry. Cells with strong signals (top 25%, TOP) and weak signals (bottom 25%, BOT) were selected. The sgRNAs from these cells were then sequenced and analyzed. For cell viability screening, cells were treated with or without PARGi for 21 days before collection. b, Scatter plot of DrugZ scores of PARG KO cells treated with or without PARGi treatment. The genes in the same pathway were marked with specific colors. A positive score indicates an enhanced pADPr signal, while a minus score indicates a decreased pADPr signal. c and d, Ranking of PARGi co-essential genes on the basis of a DrugZ analysis of the results of CRISPR/Cas9 screens performed with Toronto Knock Out Library (version 3) in HEK293A PARG KO cells and HEK239A cells. e, Analysis of biological processes of PARGi co-essential genes identified in HEK293A PARG KO cells and HEK239A cells. f, HEK293A WT cells, PARG KO, POLB KO, LIG1 KO, XRCC1 KO, and LIG3 knockdown cells were treated with different doses of PARGi for 72 hours. Cell viability was determined by the CellTiter-Glo assay. g, The cell Viability of HEK293A WT and PARG cells under POLB knockdown to PARGi. Cells were treated with different doses of PARGi for 72 hours.

Added in the revised version: As for FACS-based CRISPR screening, we showed that PARP1 consistently appeared as the gene that, when depleted, led to reduced pADPr in all settings (Figure 5B, Figure 5-figure supplement 1A, Supplementary File 1b-d), which agrees with the dominant role of PARP1 in promoting PARylation in cells. The depleted genes that led to increased pADPr signaling were also similar in all of these settings (Figure 5B, Figure 5-figure supplement 1A); these include BER genes (i.e. POLB, XRCC1, LIG3, and CHD1L) and genes involved in DNA replication/Okazaki fragment maturation (i.e. RFCs, FEN1, and LIG1). Interestingly, ARH3 only appeared in PARG KO cells treated with PARGi (Figure 5B, Figure 5-figure supplement 1A), indicating that as anticipated ARH3 may serve as a backup dePARylating enzyme.

Added in the revised version: We further investigated whether the aforementioned genes that regulated pADPr signaling would show strong synthetic lethality with PARGi in both PARG KO cells (Figure 5C, Supplementary File 1e) and WT cells (Figure 5D, Supplementary File 1f). To identify the most confident hits, we used a stringent false discovery rate (FDR) threshold of 0.0555. Consistently, PARP1 was listed as the most resistant gene with PARGi in PARG KO cells (Figure 5C). Interestingly, ARH3 appeared as the top gene in PARG KO cells treated with PARGi (Figure 5C), while other dePARylation enzymes (e.g. MARCROD1/2 andTARG1/OARD1) did not show any synthetic lethal effect in these screens. These data agree with the results of FACS-based screens and suggest that ARH3 is a backup enzyme in the absence of PARG following PARGi treatment. Indeed, loss of ARH3 further sensitized PARG KO cells to PARGi (Figure 5-figure supplement 1B). In addition, HPF1, another gene/protein involved in PARylation regulation and licensing serine mono-ADP-ribosylation by PARP1/256, was among the top synthetic lethality genes in PARG KO cells treated with PARGi, but not in WT cells (Figure 5C, Figure 5D, Figure 5-figure supplement 1C).

Added in the revised version: Thus, our CRISPR screens revealed that genes involved in PARylation regulation (ARH3 and HPF1), BER (i.e. POLB, XRCC1, LIG3, and CHD1L), and DNA replication/Okazaki fragment maturation (i.e. RFCs, FEN1, and LIG1) are responsible for pADPr signaling and/or cell lethality in WT and PARG KO cells.

Added in the revised version: To further explore the involvement of these aforementioned genes, we compared synthetic lethality genes in PARG KO cells with those in WT cells treated with PARGi. Interestingly, two BER key genes, POLB and XRCC1, showed strong synthetic lethality with PARGi in both PARG KO cells (Figure 5C) and WT cells (Figure 5D), which is consistent with the results of a previous report that POLB-deficient cells are sensitive to PARGi33. Moreover, the GO analysis revealed that the BER pathway displays synthetic lethality in both PARG KO cells and WT cells following PARGi treatment (Figure 5E). Importantly, LIG1 and LIG3, the two ligases that are directly involved in Okazaki fragment ligation, were also listed as synthetic lethality genes (Figure 5C, Figure 5D). To further validate these data, we created several KO/KD cell lines, including KO of BER genes (i.e. POLB and XRCC1) and genes involved in Okazaki fragment ligation (i.e. LIG1 and LIG3) (Figure 5-figure supplement 1D). Consistent with previous reports, XRCC1 loss destabilizes LIG3 (Figure 5-figure supplement 1D)57, 58. Notably, loss of genes involved in Okazaki fragment ligation induced S phase pADPr signaling, just like PARG KO cells following PARGi treatment, while loss of BER genes resulted in pADPr signaling throughout the cell cycle (Figure 5-figure supplement 1E). These results agree with the aforementioned results that unligated Okazaki fragments are the likely source of SSBs, which induces S phase-specific pADPr signaling. Also consistent with CRISPR screening data, loss of genes involved in Okazaki fragment ligation or BER resulted in increased sensitivity to PARGi (Figure 5E). Furthermore, we knocked down POLB in WT and PARG KO cells (Figure 5-figure supplement 1F), which resulted in increased sensitivity to PARGi in both WT and PARG KO cells (Figure 5G). Together, our data indicate that unligated Okazaki fragments induce S phase-specific pADPr, while BER intermediates lead to pADPr throughout the cell cycle; deficiency in either of these events would increase cytotoxicity following PARGi treatment.

Added in the revised version: PARG expression is a potential biomarker for PARGi-induced cytotoxicity

Added in the revised version: We showed that cells with no detectable full-length PARG expression are extremely sensitive to PARG inhibition (Figure 1, Figure 1-figure supplement 1). Thus, we are interested in determining whether PARG expression would correlate with PARGi sensitivity. An early study examined a panel of ovarian cancer cell lines and reported PARGi sensitivity in a subset of these cell lines34. We compared PARG expression in PARGi sensitive and resistant cell lines and found that PARG expression was significantly reduced in PARGi sensitive cell lines (Figure 6A). To support this result, we knocked down PARG by shRNA in HeLa cells and showed that loss or further reduction of PARG led to dramatic sensitivity to PARGi (Figure 6B). To further test our hypothesis, we ranked the PARG mRNA level in several ovarian cancer cell lines using the CCLE database (Figure 6C). We identified an ovarian cancer cell line, RMUGS, which has low PARG expression, similar to that observed in another known PARGi sensitive cell line, KURAMOCHI. Notably, RMUGS cells were considered as potential PARGi-resistant cells based on an earlier report34. However, we showed that RMUGS cells were very sensitive to PARGi treatment (Figure 6D), and the sensitivity of these cells to PARGi could be reversed by PARPi treatment (Figure 6E). These data indicate that PARG expression may dictate the sensitivity of tumor cells to PARGi-based therapy.

Added in the revised version: PARG expression is a potential marker for PARGi sensitivity.

Added in the revised version: a, PARG mRNA level comparison between sensitive (RMGI, KURAMOCHI, and OVMANA) and resistant cells (COV362, COV318, OV56, OVISE, OVSAHO, CAOV3, and OVCAR3) from Pillay’s work on the basis of CCLE data. b, Clonogenic assay results of control and PARG knockdown HeLa cells with PARGi (2 µM) treatment for 7 days. shRNA knockdown efficiency was confirmed by an immunoblot of PARG. c, Ranked PARG expression level in ovarian cancer cell lines based on the CCLE database. The sensitive and resistant cells from Pillay’s work were labeled. d, Clonogenic assay results of OVCAR3 and RMUGS treated with or without PARGi (2 µM). e, Clonogenic assay results of RMUGS treated with PARGi (2 µM), PARPi (2 µM), or both. f, Left: Representative images of PARG IHC staining in breast and ovarian tissues and tumor samples to determine the PARG expression level. The summary is listed at the bottom. Scale bar, 200 µm. Right: The scatter plot of mean immunostaining intensity of PARG in each sample. The mean of each group was plotted.

Added in the revised version: An early study reported that loss of PARG expression resulted in PARPi resistance50. They demonstrated that PARG is frequently lost in acquired PARPi-resistant mouse mammary tumors and further revealed that PARG depletion occurs in triple-negative breast and ovarian cancer. Another study showed that 60 of 274 (22%) of human ovarian tumors have low PARG expression33. Thus, PARGi can be potentially used to treat these PARPi-resistant cancers, especially those with low PARG expression. We further performed an IHC assay to detect PARG expression in breast and ovarian cancer TMA samples. The IHC assay was established with the use of our xenograft tumors derived from HeLa and HeLa PARG KO cells (data not shown). As shown in Figure 6F, PARG expression was detected in normal ovarian and breast tissues. PARG expression was also detected in a majority or most of ovarian and breast tumor samples (Figure 6F). Interestingly, about 5% to 40% of these breast or ovarian cancer tumor samples showed low expression levels of PARG, some even showed no detectable level of PARG, compared with those in normal tissue samples (Figure 6F). These data suggest that PARG downregulation occurs in breast and ovarian cancers. PARGi may be a promising strategy for the treatment of these cancers with low PARG expression.

Added in the revised version: In early studies, HR-deficient cells showed increased sensitivity to PARG inhibition18, 25, 26, 59. Similarly, we observed PARGi-induced cytotoxicity in RPE1 Flag-Cas9 TP53/BRCA1 DKO cells, but not in control cells (Figure 6-figure supplement 1A). Moreover, combining PARPi and PARGi did not reveal any additive effect (Figure 6-figure supplement 1A). Furthermore, we created inducible BRCA1 depletion cells using the auxin-inducible degron (mAID) tag. As shown in Figure 6-figure supplement 1B, HR-proficient cells showed normal sensitivity to PARGi, while HR deficiency by the loss of BRCA1 conferred modest sensitivity to PARGi. The PARGi sensitivity in BRCA1 depletion cells was reversed to normal by the restoration of HR due to 53BP1 loss. Consistent with the results obtained in RPE1 cells, the PARPi+PARGi combination did not lead to any further change in cytotoxicity. While these results indicate that PARGi can be used to target HR-deficient cancers, the effect is quite modest. Moreover, PARGi treatment did not show any effect on PARPi-resistant cells due to restoration of HR mediated by 53BP1 loss.

Added in the revised version: To further investigate the potential contribution of HR deficiency and/or PARG loss to PARGi sensitivity, we knocked down PARG by shRNA in these inducible BRCA1 depletion cells with or without 53BP1 KO (Figure 6-figure supplement 2A). As expected, PARG loss is the main driver of PARGi sensitivity; regardless whether these cells are HR-proficient, HR-deficient, or HR-restored because of 53BP1 loss (Figure 6-figure supplement 2B). Our data therefore strongly suggest that PARG expression is the major determinant of PARGi sensitivity.

Added in the revised version: PARG is essential for cell survival

Added in the revised version: Our results suggested that cells with no detectable full-length PARG expression or low expression of PARG are extremely sensitive to PARG inhibition (Figure 1, Figure 6). As aforementioned, the question is why PARG KO cells would be sensitive to PARGi. To address this question, we first performed whole-genome CRISPR screening. However, the potential second target was not revealed by our CRISPR screening data (Figure 5). The second possibility is that our PARG KO cells are not complete PARG KO cells. These cells may have residual PARG expression or activity and only cells with very low PARG expression are sensitive to PARGi, as shown in Figure 6. Indeed, two independent Parg knockout mice have been created by targeting early exons of PARG13, 60. While one group of Parg knockout mice showed embryonic lethality and cells derived from these mice only survived in the presence of PARPi13, the other Parg knockout mice did not display any major phenotypes, probably because truncated Parg and Parg activity were still detected60. We favor the second possibility, especially since the PARG catalytic domain is at its C-terminus.

Added in the revised version: To test this possibility, we first used two independent antibodies that recognize the C-terminus of PARG to detect potential residual PARG isoforms or spliced variants (Figure 7-figure supplement 2A). Unfortunately, besides full-length PARG, these antibodies also recognized several other bands, some of them were reduced or absent in PARG KO cells, others were not. Thus, we could not draw a clear conclusion which functional isoform/truncated form was expressed in our PARG KO cells. Then, we used two additional gRNAs which target respectively the beginning of the catalytic domain and the sequence around the catalytic site (Figure 7A). Interestingly, we were only able to obtain viable clones with no detectable PARG in the presence of PARPi in either 293A or HeLa cells, which we named as PARG complete/conditional KO cells (cKO) (Figure 7B). Indeed, these cKO cells could not survive without PARPi (Figure 7C). To further investigate whether cell lethality is associated with PARG activity in these cells, we reconstituted PARG cKO cells with WT PARG or a PARG catalytic-inactive mutant. Consistent with previous result (Figure 1, Figure 1-figure supplement 1), WT PARG not only reduced pADPr level but also rescued cell lethality due to PARG loss, while the catalytic-inactive mutant PARG failed to do so (Figure 7D).

Added in the revised version: PARG is essential for cell survival.

Added in the revised version: a, Diagram of full-length PARG was presented with the indicated gRNAs (gRNA#3 and 4) which target different regions in the C-terminal catalytic domain. The boundary of the catalytic domain was depicted based on Uniprot annotation. gRNA#1 and gRNA#2 were used previously to generate the aforementioned HEK293A- and HeLa-derived PARG KO cells, respectively; while gRNA#3 and gRNA#4 were used to generate PARG complete/conditional knockout (cKO) in the presence of olaparib in HEK293A and HeLa cells. b, Immunoblotting was conducted to confirm the loss of PARG in PARG cKO cells derived from HEK293A and HeLa cells, which were cultured in the presence of 100 nM olaparib. c, Clonogenic assay results of WT and PARG cKO cells treated with or without PARPi (100nM) for 7 days. d, Left: The immunoblots to confirm reconstitution with WT PARG or catalytic inactivation PARG in HEK293A PARG cKO cells. Right: Results of clonogenic survival assay with HEK293A PARG cKO cells reconstituted with WT or catalytic inactivation mutant of PARG for 7 days. e, Representative clonogenic results conducted in HEK293A PARG cKO cells treated with NAM (100 µM) or NMN (1 mM) for 7 days. f, HEK293A PARG cKO cells were synchronized with double thymidine block (DTB). Cells remained with DTB or were released from DTB for 4 hours, and then fixed and stained with anti-pADPr antibody and FxCycle Violet dye. g, Immunoblots of soluble and chromatin-bound PARP1 and pADPr levels in HEK293A WT and PARG cKO cells treated with DMSO or olaparib (10 µM) for 2 hours.

Added in the revised version: The data above indicate that our PARG KO cells have residual PARG activity and treatment with PARGi in these cells further decreases PARG activity, which mimic the complete loss of PARG as that in PARG cKO cells. To further test this hypothesis, we measured the relative PARG dePARylation activity in whole cell lysates (WCL) prepared from different cell lines (Figure 7-figure supplement 2C). Consistent with our working hypothesis, residual PARG activity was detected in WCL prepared from PARG KO cells, while PARG activity was barely detected or absent in WCL isolated from PARG cKO cells. Please note that we incubated cell lysates with substrates overnight to evaluate the maximum level of pADPr hydrolysis, i.e. PARG activity, we were able to detect in these assays. It is very likely that the PARG activity in PARG KO cells was much lower than that indicated in Figure 7-figure supplement 2C, due to saturation of signals for lysates isolated from wild-type cells. Thus, the data presented here may under-estimate the reduction of PARG activity in PARG KO cells. Nevertheless, these data indicate the residual PARG activity in PARG KO cells, which is absent in PARG cKO cells. Furthermore, the PARG activity was further inhibited by PARGi in a dose-dependent manner in WCLs prepared from both WT and PARG KO cells (Figure 7-figure supplement 2D, Figure 7-figure supplement 2E). Notably, the PARG activity was consistently lower in WCL prepared from PARG KO cells when compared with that in control wild-type cells in the presence of a wide range of PARGi concentrations. However, a significant fraction of PARG activity still existed in WCL prepared from wild-type cells even at the highest concentration of PARGi used in this study (Figure 7-figure supplement 2E). These data indicate residual PARG activity in PARG KO cells likely accounts for the survival of these cells. Moreover, further inhibition of this activity by PARGi likely leads to cell lethality.

Added in the revised version: We showed above that NAD+ precursors (NAM and NMN) were able to rescue cytotoxicity of PARGi in PARG KO cells, probably by NAD+ depletion and also that these precuresors can also inhibit PARP1/2 activities. Consistently, PARG cKO cells could survive in the presence of NAD+ precursors (NAM and NMN) (Figure 7E). Taken together, PARG is an essential gene and our PARG KO cells treated with PARGi can mimic the complete loss of PARG activity. These data agree with the genetically engineered mouse models as we discussed above13.

Added in the revised version: We showed that S phase pADPr signals were specifically detected in PARG KO cells treated with PARGi, which likely originated from unligated Okazaki fragments (Figure 3). As expectedly, we observed S phase-specific pADPr in PARG cKO cells under double-thymidine block and release condition (Figure 7F). Moreover, the levels of chromatin-bound PARP1 and pADPr signal, which reflects PARylated PARP1 and/or other PARP1 substrates, increased in PARG cKO cells (Figure 7G). Treatment with PARPi decreased pADPr signals but increased chromatin-bound PARP1 (Figure 7G), which is consistent with the results obtained in PARG KO cells (Figure 4A). Together, PARG, as an essential gene, suppresses initial S phase pADPr signaling and PARP1 chromatin accumulation, which eventually lead to cell death caused by uncontrolled pADPr accumulation throughout cell cycle.

Added in the revised version: Discussion

Added in the revised version: Aberrant accumulation of pADPr and/or NAD+ exhaustion disrupts multiple cell processes, including DNA repair, replication stress response, and transcription regulation; therefore, it is highly toxic and eventually leads to cell death10, 22, 38, 39, 61. Although ADP-ribose hydrolases, especially PARG, have a protective effect against excessive PARP1 engagement, it remains unclear when and how PARG activity is engaged during normal cell proliferation. In this study, we showed that the major function of PARG is to regulate S phase-specific pADPr by PARP1, which is activated by unligated Okazaki fragments. When PARG activity is inhibited, these uncontrolled pADPr would eventually result in DNA damage and cell death. Using unbiased FACS- and cell viability-based genome-wide CRISPR screens, we uncovered additional pADPr modulators, which include the genes involved in BER (POLB, XRCC1, and LIG3), Okazaki maturation (LIG1 and FEN1), ARH3, and HPF1. More importantly, we determined that PARG expression is a critical biomarker for PARGi sensitivity. Furthermore, we showed that PARG is an essential gene by suppressing pADPr accumulation. Taken together, our study uncovers a normal function of PARG in proliferating cells: it removes pADPr generated by PARP1 at unligated Okazaki fragments during DNA replication. Of course, this normal function of PARG is also required under conditions of excessive PARP1 engagement because of defects in Okazaki fragment maturation and/or BER.

Added in the revised version: An important lesson we learned from this study is that it is challenging to confirm the complete knockout of a gene of interest. As we reported here, we initially used two independent sgRNAs, which in theory target all active forms of PARG, and created KOs in two cell lines. Additionally, we validated our KO clones by Western blotting, DNA sequencing and MMS-induced PARylation. Despite these efforts and our inability to detect full-length PARG in our KO clones, these PARG KO cells still express one or more active fragments of PARG, probably due to alternative splicing and/or alternative ATG usage. In this study, we were able to measure PARG activity directly in cell lysates using a biochemical assay, which supports our working hypothesis. However, such assay is not available for many enzymes or proteins without known enzymatic activities. Thus, one should be cautious when drawing conclusions based on KOs, especially conclusions such as genetic interactions that rely on multiple KOs.

Added in the revised version: PARylation is a reversible post-translational modification. In this study, we showed that timely and efficient dePARylation is critical for cell survival. Thus, both PARylation and dePARylation play important functions in the cell. Particularly, dePARylation appears to be indispensable, since it is required for cell survival. The successful clinical application of PARylation inhibition, such as PARPis for the treatment of BRCA-deficient cancers, raises the possibility of targeting dePARylation as an alternative strategy, especially for the treatment of PARPi-resistant tumors. Given that PARG accounts for more than 90% of dePARylation activity7, 29, selective PARG inhibitors are being developed as potential anti-cancer agents. Although defective HR and DDR have been shown to be synthetic lethal with PARG inhibition, comprehensive investigations of cellular response to PARG inhibition have not yet been conducted. In this study, we showed that unexpectedly, PARG is the top synthetic lethality gene in the CRISPR screening with our homemade DDR sgRNA library, which we further validated in HEK293A and HeLa PARG KO cells (Figure 1B, Figure 1-figure supplement 1D). Our PARG KO cells plus PARGi mimic complete loss of PARG activity (Figure 7-figure supplement 2), which provides is an unique opportunity to study the major function of PARG in proliferating cells. These data were further confirmed with the use of PARG cKO cells we generated (Figure 7).

Added in the revised version: PARP1 plays important roles during normal S phase, especially at DNA replication forks and two-ended DSBs43. More recently, unligated Okazaki fragments were identified as the source of SSBs that activate PARP1 in normal proliferating cells6. Moreover, unligated Okazaiki fragments, on which PARP1 may be trapped, can be converted into structures that result in synthetic lethality in HR-defective cancer cells62. PARG is also enriched at DNA replication forks through pADPr- and PCNA-dependent mechanisms14, 63, 64. In addition, PARG may protect DNA replication forks during normal and damage-treated S-phase15, 31, 65. Interestingly, the identification of unligated Okazaki fragments as the source of endougenous pADPr in normal proliferating cells was only noticeable under short-term PARG inhibition. Treatment with another recently developed PARG inhibitor COH34, also led to increased PARylation at replication sites in S phase cells26. These previous reports agree with the data presented in this study. Indeed, we showed robust and specific PARP1/2-dependent S phase pADPr in PARG KO cells treated with PARGi (Figure 3A, Figure 3-figure supplement 1). We further confirmed that the S phase-specific pADPr requires normal S phase progression and the generation of Okazaki fragments (Figure 3B-C, Figure 3-figure supplement 1B). Moreover, the cytotoxicity of PARGi requires S phase progression (Figure 3D-E). In addition, PARG KO cells displayed more chromatin-bound PARP1 and PARylated PARP1 (Figure 4A, Figure 4-figure supplement 1A). Enhanced pADPr levels were observed in the chromatin fraction of PARG KO cells under PARGi treatment (Figure 4A). One possible explanation is that in PARG KO cells treated with PARGi, chromatin-bound or trapped PARP1 may not be removed appropriately or efficiently, which further prevents Okazaki fragment ligation/maturation that eventually results in cell lethality. However, the mechanisms underlying chromatin accumulation of PARP1 and PARylated PARP1 in PARG KO cells remain to be further elucidated.

Added in the revised version: Additionally, we do not yet know all the substrates of PARG in S phase cells. We speculate that PARP1 is one of the major PARG substrates in S phase cells. As mentioned above, chromatin-bound PARP1 as well as PARylated PARP1 increased in PARG KO cells. Moreover, PARP1 depletion was able to rescue cell lethality in PARG cKO cells or PARG KO cells treated with PARGi. Of course, PARG may have additional substrates besides PARP1 which are required for its roles in S phase progression. Precisely how PARG regulates S phase progression warrants further investigation.

Added in the revised version: The inability to remove S phase-specific pADPr eventually led to pADPr throughout cell cycle and γH2AX signaling (Figure 4C-D, Figure 4-figure supplement 1D-F), indicating that it is essential to resolve S phase-specific pADPr to maintain cell viability. In agreement with the literature6, 26, we observed S phase-specific pADPr in HeLa cells following PARGi treatment, although such S phase-specific pADPr was not detected in HEK293A cells following PARGi treatment. The lower level of PARG expression may be the reason for this difference between HEK293A and HeLa cells (Figure 4-figure supplement 1E). However, significant amounts of PARG activity still existed in HeLa cells, since these cells did not show noticeable PARGi sensitivity unless we knockout or knockdown PARG.

Added in the revised version: Cell viability–based genome-wide CRISPR/Cas9 screening unbiasedly uncovers sensitivity and resistance genes under certain conditions49, while FACS-based CRISPR/Cas9 screening using reporter cell lines or antibodies recognizing specific signaling molecules can reveal key regulators of signaling pathways53. In this study, we performed both types of screens, which revealed key pADPr regulators, including BER pathway proteins (POLB, XRCC1, LIG3, and CHD1L), DNA replication/Okazaki fragment maturation proteins (RFCs, FEN1, and LIG1), and ARH3 (Figure 5B, Figure 5-figure supplement 1A). In addition, consistent with our hypothesis that pADPr accumulation leads to cell death, we showed that loss of genes involved in Okazaki fragment maturation or BER resulted in PARGi sensitivity (Figure 5C-G). All of these results further validated our working model that S phase-specific pADPr likely originated from unligated Okazaki fragments, which if not removed would eventually lead to DNA damage and cell death.

Added in the revised version: During DNA replication, Okazaki fragments are normally ligated by LIG1. However, some unligated Okazaki fragments may activate PARP1 and recruit XRCC1/LIG36, which agrees with the results of early studies indicating that both LIG1 and LIG3 are required for Okazaki fragment maturation42. Indeed, we showed that LIG1 depletion, which would result in more DNA nicks in S phase cells, led to the enhancement of S phase PARylation (Figure 5-figure supplement 1E). As for LIG3, it can substitute LIG1 function for the ligation of Okazaki fragments42, 66. However, it also has a critical and essential function in mitochondria but is separated for its role in XRCC1-dependent SSBR67. We observed minor S phase pADPr due to LIG3 knockdown with PARGi treatment (Figure 5-figure supplement 1E), which is likely due to insufficient KD of LIG3 in these experiments. Unlike LIG1 KO and LIG3 KD cells, XRCC1 KO and POLB KO cells increase PARylation in a cell cycle–independent manner (Figure 5-figure supplement 1E), which agrees with the finding that POLB and XRCC1 are required for BER/SSB repair, but not specifically for Okazaki fragment ligation/maturation. Our results suggest that PARylation and dePARylation are mainly involved in two cellular processes, i.e. Okazaki fragment maturation and BER/SSB repair, since both processes have nicked DNA as intermediates, which are likely the physiological substrates that recruit and activate PARP1/2.

Added in the revised version: Interestingly, two well-known PARylation regulators, ARH3 and HPF1, showed synthetic lethality with PARGi in PARG KO cells (Figure 5C). Given that ARH3 and PARG are the primary dePARylation enzymes10, it should be predictable that ARH3 was listed as the top hit with PARGi in PARG KO cells in both pADPr signal screen and synthetic lethality screen (Figure 5B-C). As for HPF1, a previous report showed that the loss of HPF1 will release the activity of PARP1 to PARylation in acidic resides56; moreover, a recent study shows that HPF1 and nucleosomes mediate a dramatic switch in the activity of PARP1 from polymerase to hydrolase68. More important, HPF1 promotes the ligation of the Okazaki fragment by LIG3-XRCC1 as a backup pathway69, while the unligated Okazaki fragments were the source of the S phase pADPr signal in this study.

Added in the revised version: More importantly, our data indicate that PARG expression is a potential biomarker for PARGi sensitivity. We used the available database and cell lines to validate that cells with low PARG expression are sensitive to PARGi. In addition, we examined PARG expression by the IHC assay in commercial TMA samples and uncovered a fraction of breast and ovarian cancers with no detectable or low PARG expression (Figure 6). We also examined the sensitivity of HR-deficient cells to PARGi. Although HR deficiency sensitizes cells to PARGi, our results indicate that PARG loss or low expression is still the main driver of PARGi sensitivity in these cells (Figure 6-figure supplement 1-2).

Added in the revised version: Mechanically, we showed that the dePARylation activity of PARG is indispensable for cell survival (Figure 7, Figure 7-figure supplement 2). The residual PARG dePARylation activity observed in PARG KO cells likely supports cell growth, which can be further inhibited by PARGi (Figure 7-figure supplement 2). Although a dose-dependent inhibition of PARG activity by PARGi was also noted in wild-type cells, significant fraction of PARG activity remained even in the presence of highest concentration of PARGi used in this study, which is consistent with our results that HEK293A cells are insensitive to PARGi. More importantly, these data further indicate that PARG expression/activity is a potential biomarker for PARGi sensitivity.

Added in the revised version: Taken together, the dramatic S phase-specific pADPr signaling detected under our experimental conditions provides us a rare opportunity to gain a better understanding of PARG function and pADPr signaling in normal proliferating cells without any exogenous DNA damage. In addition, our systematic CRISPR/Cas9 screens uncover key regulators of pADPr signaling that also contribute to PARGi sensitivity. Moreover, our finding that PARG expression may a potential biomarker of PARGi sensitivity will allow the further development of these inhibitors for the treatment of tumors with PARG loss and therefore offer a new targeted strategy for cancer therapy.

Materials and Methods

Cell culture

Earlier version: HEK293A, HeLa, and OVCAR3 cells were purchased from the American Type Culture Collection (ATCC); HEK293A and HeLa cells were cultured in DMEM with 10% fetal bovine serum, while OVCAR3 cells were cultured in RPMI with 10% fetal bovine serum. RPE1-hTERT FLAG-Cas9 TP53−/− BRCA1−/− cells were a gift that was kindly provided by Dr. Daniel Durocher (University of Toronto) and were cultured in DMEM with 10% fetal bovine serum. RMUGS was obtained from the Japanese Collection of Research Bioresources Cell Bank and cultured in Ham’s F12 medium with 10% fetal bovine serum. Knockout cells generated in 293A and HeLa cells were created with pLentiCRISPRv2 (Addgene, #52961) containing indicated gRNAs (Fig. S8), as described previously33. All knockout cells were validated by Western blotting and DNA sequencing (Fig. S8). The 293A-derived PARP1/2 DKO cells were the same as those in the previous study33. All cell lines were free of mycoplasma contamination.

Revised version: HEK293A, HeLa, and OVCAR3 cells were purchased from the American Type Culture Collection (ATCC); HEK293A and HeLa cells were cultured in DMEM with 10% fetal bovine serum, while OVCAR3 cells were cultured in RPMI with 10% fetal bovine serum. RPE1-hTERT FLAG-Cas9 TP53−/− BRCA1−/− cells were a gift that was kindly provided by Dr. Daniel Durocher (University of Toronto) and were cultured in DMEM with 10% fetal bovine serum. RMUGS was obtained from the Japanese Collection of Research Bioresources Cell Bank and cultured in Ham’s F12 medium with 10% fetal bovine serum. Knockout cells generated in 293A and HeLa cells were created with pLentiCRISPRv2 (Addgene, #52961) containing indicated gRNAs (Figure 7-figure supplement 1), as described previously70. All knockout cells were validated by Western blotting and DNA sequencing (Figure 7-figure supplement 1). The 293A-derived PARP1/2 DKO cells were the same as those in the previous study70. All cell lines were free of mycoplasma contamination.

Chemical reagents and antibodies

The NAD/NADH Assay Kit II (colorimetric) (ab221821) was obtained from Abcam. The CellTiter-Glo Luminescent Cell Viability Assay (G7573) was purchased from Promega. Nicotinamide (NAM, S1899), β-nicotinamide mononucleotide (NMN, S5259), olaparib (S1060), and FK866 (S2799) were purchased from SelleckChem. PDD 00017272 (HY-133531), PDD 00017238 (HY-133530), PDD 00031705 (HY-135846), and PDD 00017273 (HY-108360) were obtained from MedChemExpress. Methyl methanesulfonate (MMS, 129925), thymidine (T9250), emetine (E2375), and crystal violet solution (HT90132) were from Sigma. The propidium iodide (PI, P3566), GATEWAY cloning system (11789100, 11791100) and FxCycle Violet Stain (F10347) were obtained from ThermoFisher. The QuikChange II Site-Directed Mutagenesis Kit (200523) was from Agilent. The CometAssay Single Cell Gel Electrophoresis Assay (4250-050-K) was obtained from Trevigen.

Antibodies against PARP1 (9532S), PARG (66564S, which mainly recognizes full-length isoform 1, but also weakly recognizes isoforms 2 and 3 (data not shown)), and XRCC1 (2735S) were purchased from Cell Signaling Technology. Antibodies against C-terminus of PARG were purchased from Santa Cruz Biotechnology (SC-398563) and Novus Biologicals (NBP2-55661). Antibodies against histone H3 (ab18521), γH2AX (ab2893), and POLB (ab26343) were from Abcam. Antibodies against β-actin (A5316), tubulin (T6199), and ARH3 (HPA027104) were from Sigma. The antibodies anti-pADPr (10H) (sc-56198), LIG1 (sc-56087), and BRCA1 (sc-6954) were from Santa Cruz Biotechnology. Anti-PARP2 antibody (39743) was from Creative Motif, anti-LIG3 antibody (GTX70143) was from GeneTex, and anti-53BP1 antibody (NB100-304H) was from Novus Biologicals. Anti-γH2AX (05-6361) and PAR (10H) (AM80) antibodies were obtained from Millipore. Alexa Fluor Plus 488 secondary antibody goat anti-mouse (A327230) was from ThermoFisher.

Plasmid constructs

Earlier version: PARP1 (HsCD00043719) and PARG (HsCD00859023) cDNAs, purchased from DNASU, were subjected to mutagenesis using the QuikChange II Site-Directed Mutagenesis Kit. WT and mutant constructs were subsequently cloned into the vector pLenti CMV Neo DEST (705-1) (Addgene, #17392) by the GATEWAY cloning system. shRNA constructs (pGIPZ-based vector) targeting PARG (Clone ID: V2LHS_11965, V2LHS_250448, V3LHS_379835, and V3LHS_379832) and targeting POLB (Clone ID: V2LHS_170201 and V2LHS_222459) were obtained from Horizon Discovery Biosciences. The Toronto KnockOut Library v3 (TKOv3) (90294) was from Addgene. The DNA Damage Response MKOv4 Library (Addgene, #140219) was from our previous study34.

Revised version: PARP1 (HsCD00043719) and PARG (HsCD00859023) cDNAs, purchased from DNASU, were subjected to mutagenesis using the QuikChange II Site-Directed Mutagenesis Kit. WT and mutant constructs were subsequently cloned into the vector pLenti CMV Neo DEST (705-1) (Addgene, #17392) by the GATEWAY cloning system. shRNA constructs (pGIPZ-based vector) targeting PARG (Clone ID: V2LHS_11965, V2LHS_250448, V3LHS_379835, and V3LHS_379832) and targeting POLB (Clone ID: V2LHS_170201 and V2LHS_222459) were obtained from Horizon Discovery Biosciences. The Toronto KnockOut Library v3 (TKOv3) (90294) was from Addgene. The DNA Damage Response MKOv4 Library (Addgene, #140219) was from our previous study35.

Cell viability assays

To measure cell viability, ∼2,000 cells were seeded into the 96-well plates. Drugs with indicated concentrations (three biological replicates) were added after 24 hours. After 72 hours of treatment, cell numbers were measured by the CellTiter-Glo Luminescent Cell Viability Assay according to the manufacturer’s instructions to draw the cell viability curves.

NAD+ level measurement

The NAD+ relative level was measured following the manufacturer’s instructions. In brief, the same number of cells with the indicated treatment were collected and lysed with NAD+ extraction solution (0.5 M perchloric acid) for 30 min on ice. The extraction was neutralized by 0.55 M K2CO3 and centrifuged. The supernatant was collected and added to a 96-well plate. NAD+ standard and the NAD+ reaction mixture were added to the 96-well plate for the next colorimetric reading at OD 450 nm after a 30-min reaction. For each time point, at least three biological replicates were measured.

Clonogenic assay

To assess cellular sensitivity to PARGi or other agents, the indicated cells were seeded on 12-well plates and subsequently exposed to DMSO or indicated treatments for 7 to 14 days. After phosphate-buffered saline (PBS) washes, cells were stained with crystal violet solution. For each condition, three biological replicates were included and the representative results were presented.

Western blotting analysis

Unless further operation was indicated, such as chromatin and soluble fractionation, cells were washed with PBS and directly lysed on the plate by 2× Laemmli buffer, boiled at 95 °C for 10 min, and separated by sodium dodecyl sulfate polyacrylamide gel electrophoresis, transferred to membranes, and immunoblotted with the indicated antibodies. For each condition, at least two biological replicates were conducted and the representative results were presented.

Fluorescence-activated single cell sorting analysis (FACS)

The same number of trypsinized cells with the indicated treatments were collected into a 15-ml tube with PBS, which were subsequently fixed by pre-cold ethanol to a final 70% concentration. After centrifugation, cell pellets were permeabilized by PBS with 0.5% Triton X-100 at ambient temperature for 20 min. After being washed with PBS, cell pellets were blocked by PBS with 4% bovine serum albumin at ambient temperature for 1 hour. The cell pellets were incubated at 4 °C overnight with antibodies (pADPr, AM80; γH2AX, 05-6361) at 1:1000 dilution by PBS with 4% BSA. After being washed with PBS three times, cell pellets were incubated with an Alexa Fluor 488 secondary antibody at 1:500 dilution for another 1 hour at ambient temperature. The pellets were then stained with FxCycle Violet Stain (1 μg/ml) or PI (20 μg/ml) with RNase A (10 μg/ml) and resuspended in PBS for the next flow cytometry analysis by a BD C6 flow cytometer (Becton Dickinson) or Attune Flow cytometers (ThermoFisher). FlowJo software (v10.6.1) was used to analyze the acquired data. For each condition, at least three biological replicates were conducted and the representative results were presented.

Chromatin and soluble fractionation and enrichment for PARylated proteins

Earlier version: Cells with the indicated treatments were collected for chromatin and soluble fractionation with a protocol reported previously35, 36. In brief, cells were lysed in NETN buffer (20 mM Tris-HCl [pH 8.0], 1 mM EDTA, 100 mM NaCl, 0.5% NP-40, and 1 mM DTT) containing proteinase inhibitors for 20 min on ice. After centrifugation, the soluble fraction (i.e. the supernatant) was collected into fresh tubes. The chromatin fractionation (i.e. the pellets) was washed twice with NETN buffer. The soluble and chromatin fraction were both diluted with 2× Laemmli buffer, boiled at 95 °C for 10 min, and subjected to Western blotting analysis. Three biological replicates of chromatin and soluble fractionation were conducted and the representative results were presented.

Revised version: Cells with the indicated treatments were collected for chromatin and soluble fractionation with a protocol reported previously48, 49. In brief, cells were lysed in NETN buffer (20 mM Tris-HCl [pH 8.0], 1 mM EDTA, 100 mM NaCl, 0.5% NP-40, and 1 mM DTT) containing proteinase inhibitors for 20 min on ice. After centrifugation, the soluble fraction (i.e. the supernatant) was collected into fresh tubes. The chromatin fractionation (i.e. the pellets) was washed twice with NETN buffer. The soluble and chromatin fraction were both diluted with 2× Laemmli buffer, boiled at 95 °C for 10 min, and subjected to Western blotting analysis. Three biological replicates of chromatin and soluble fractionation were conducted and the representative results were presented.

For enrichment of PARylated proteins in chromatin and soluble fractions, cells were lysed for 20 min on ice by NETN buffer containing proteinase inhibitors and PARP/PARGis (10 μM) to avoid PARylation/dePARylation during cell lysis. After fractionation, the chromatin fractionation (i.e. the pellets) was resuspended by NETN buffer containing proteinase inhibitors and PARP/PARGis (10 μM) and subjected to sonication and centrifugation. The chromatin and soluble fractions were incubated with Af1521 macrodomain affinity resin for 4 hours at 4 °C. After being washed three times with NETN buffer, the beads were boiled with 1× Laemmli buffer at 95 °C for 10 min. The elution was analyzed by Western blotting with the indicated antibodies. Two biological replicates of enrichment were conducted and the representative results were presented.

shRNA knockdown

pGIPZ vectors containing scrambled or gene-specific shRNAs were packed into lenti-virus with the packaging vector psPAX2, the envelope vector pMD2.G, and polyethyleneimine; then lenti-virus was used to infect indicated cell lines. Infected cells were selected with puromycin for ∼5 days. The pooled cells were validated by Western blotting and used for further experiments.

Cell viability–based and flow cytometry–based CRISPR screens

Earlier version: DDR library screening was conducted as described previously34. In brief, cells were infected with DDR library virus at a low multiplicity of infection (MOI) (~ 0.3) for 24 hours and then selected with puromycin (2 μg/ml) for 72 hours. Cells were collected as the initial time point T0, and the remaining cells (five million cells in each replicate, each condition containing at least two replicates) were passaged every 3 days for 21 days. Five million cells were collected at both T0 and the final time point at 21 days (T21).

Revised version: DDR library screening was conducted as described previously35. In brief, cells were infected with DDR library virus at a low multiplicity of infection (MOI) (∼ 0.3) for 24 hours and then selected with puromycin (2 μg/ml) for 72 hours. Cells were collected as the initial time point T0, and the remaining cells (five million cells in each replicate, each condition containing at least two replicates) were passaged every 3 days for 21 days. Five million cells were collected at both T0 and the final time point at 21 days (T21).

Earlier version: For whole-genome CRISPR gRNA screening, 120 million cells were infected with lentiviruses encoding the TKOv3 library at a low MOI ratio (< 0.3) for 24 hours, as conducted previously36. Infected cells were selected with puromycin (2 μg/ml) for 72 hours. For the cell viability–based screens, 20 million cells were collected after selection and marked as the initial time point (T0). The remaining cells (20 million cells in each replicate, with each condition containing at least two replicates) were passaged every 3 days for 21 days and treated with PARGi or DMSO. Twenty million cells were collected in each condition after 21 days and marked as T21. For flow cytometry-based screening, a similar workflow was performed as described previously37. At day 5 after selection, 150 million cells in each replicate were treated with either PARGi (10 μM) for 4 hours or DMSO. Cells were subjected to sequential fixation, permeabilization, blocking, antibody incubation, and staining, as described above. After that, ~ 20 million cells were collected in both the top 25% with the highest signal and the bottom 25% with the lowest signal cell populations on the basis of the pADPr signal, as determined by flow cytometry.

Revised version: For whole-genome CRISPR gRNA screening, 120 million cells were infected with lentiviruses encoding the TKOv3 library at a low MOI ratio (< 0.3) for 24 hours, as conducted previously49. Infected cells were selected with puromycin (2 μg/ml) for 72 hours. For the cell viability–based screens, 20 million cells were collected after selection and marked as the initial time point (T0). The remaining cells (20 million cells in each replicate, with each condition containing at least two replicates) were passaged every 3 days for 21 days and treated with PARGi or DMSO. Twenty million cells were collected in each condition after 21 days and marked as T21. For flow cytometry-based screening, a similar workflow was performed as described previously53. At day 5 after selection, 150 million cells in each replicate were treated with either PARGi (10 μM) for 4 hours or DMSO. Cells were subjected to sequential fixation, permeabilization, blocking, antibody incubation, and staining, as described above. After that, ∼ 20 million cells were collected in both the top 25% with the highest signal and the bottom 25% with the lowest signal cell populations on the basis of the pADPr signal, as determined by flow cytometry.

Genomic DNA was extracted from cell pellets by the QIAamp Blood Maxi Kit (Qiagen) and resuspended in Buffer EB (10 mM Tris–HCl [pH 7.5]) after precipitation by ethanol and sodium chloride. DNA was then amplified by PCR with primers harboring Illumina TruSeq adapters with i5 and i7 barcodes, and the resulting libraries were sequenced on an Illumina NextSeq 500 system. The BAGEL algorithm (https://github.com/hart-lab/bagel) was used to calculate essentiality scores. A DrugZ analysis (https://github.com/hart-lab/drugz) was used to calculate the difference between different groups.

Alkaline comet assay

An alkaline comet assay was performed using a CometAssay kit under the manufacturer’s instructions. In brief, equal numbers of cells from different treatments were collected at the same time and resuspended by PBS with a concentration of 1 × 105 cells/ml. The cell suspension (50 μl) was well-mixed with 500 μl of molten LMAgarose (at 37 °C) and then immediately spread onto CometSlide. After solidification at 4 °C in the dark for 10 min, slides were incubated with lysis solution for 60 mins at 4 °C in the dark and then immersed into the alkaline unwinding solution (200 mM NaOH, 1 mM EDTA, pH>13) for 60 min at 4 °C in the dark. After that, slides were placed into an electrophoresis slide tray with the alkaline unwinding solution for 45 min at 23 voltage. Slides were gently washed, twice with ddH2O and once with 70% ethanol. Slides were dried at 37 °C for ∼20 min to make all cells in a single plane and stained with SYBR-Gold. Images were captured by a Nikon 90i microscope at ×10 magnification and analyzed by OpenComet (V1.3.1).

CRISPR/Cas9-mediated stable knockout cells and reconstitution

Removed in the revised version: Cells were transfected with pLentiCRISPRv2 plasmids containing the gRNAs targeting indicated genes (Fig. S8) using polyethyleneimine. After 24 hours, cells were selected with puromycin for another 48 hours and seeded into 96-well plates with 1 cell in each well. After 2 weeks, the single clones were selected from 96-well plates for further validation by Western blotting and DNA sequencing (Fig. S8). Each knockout clone was generated with single gRNA.

Earlier version: pLenti CMV Neo DEST expression vectors containing empty or indicated gene constructor were packed into lenti-virus with the packaging vector psPAX2, the envelope vector pMD2.G, and polyethyleneimine. Cells were infected with the indicated virus and selected with puromycin for ~5 days, after 24 hours. The pooled cells were confirmed with Western blotting to validate the expression of genes and used in further experiments.

Revised version: Cells were transfected with pLentiCRISPRv2 plasmids containing the gRNAs targeting indicated genes (Figure 7-figure supplement 1) using polyethyleneimine. After 24 hours, cells were selected with puromycin for another 48 hours and seeded into 96-well plates with 1 cell in each well. After 2 weeks, the single clones were selected from 96-well plates for further validation by Western blotting and DNA sequencing (Figure 7-figure supplement 1). Each knockout clone was generated with single gRNA. pLenti CMV Neo DEST expression vectors containing empty or indicated gene constructor were packed into lenti-virus with the packaging vector psPAX2, the envelope vector pMD2.G, and polyethyleneimine. Cells were infected with the indicated virus and selected with puromycin for ∼5 days, after 24 hours. The pooled cells were confirmed with Western blotting to validate the expression of genes and used in further experiments.

Earlier version: The generation of PARG complete/conditiona knockout (cKO) cells

Revised version: The generation of PARG complete/conditional knockout (cKO) cells

Earlier version: Cells were transfected with pLentiCRISPRv2 plasmids containing the gRNAs targeting the indicated sequences (Fig. 7A) with polyethyleneimine. After selection with puromycin, cells were seeded into 96-well plates with 1 cell each well in the presence of low concentration of olaparib (100 nM). Each cKO clone was generated with single gRNA. Viable clones were cultured with low concentration of olaparib and subjected to further validation by Western blotting and DNA sequencing (Fig. S9). To keep PARG cKO cells viable, validated clones were always cultured with the addition of low concentration of olaparib.

Revised version: Cells were transfected with pLentiCRISPRv2 plasmids containing the gRNAs targeting the indicated sequences (Figure 7A) with polyethyleneimine. After selection with puromycin, cells were seeded into 96-well plates with 1 cell each well in the presence of low concentration of olaparib (100 nM). Each cKO clone was generated with single gRNA. Viable clones were cultured with low concentration of olaparib and subjected to further validation by Western blotting and DNA sequencing (Figure 7-figure supplement 2). To keep PARG cKO cells viable, validated clones were always cultured with the addition of low concentration of olaparib.

Double thymidine block

After cells had adhered to tissue culture plates, they were treated with 2 mM thymidine for 16 hours. The medium was removed, and the cells were washed with sterile PBS twice. Fresh medium was added for another 9 hours. After that, cells were incubated with thymidine for an additional 14 hours.

Human tissue IHC analysis

Human ovary and breast carcinoma tissue microarrays were purchased from US Biomax. The ovary carcinoma tissue (BC11115d) contains 5 cases of clear cell carcinoma, 62 serous carcinoma, 10 mucinous adenocarcinoma, 3 endometrioid adenocarcinoma, 10 lymph node metastasis carcinoma, 10 adjacent normal ovary tissue (missing one case adjacent normal ovary tissue and one case of serous carcinoma, so we have 98 cases in total for analysis). The breast carcinoma tissue (BC081120f) contains 100 cases of invasive carcinoma of no special type, 10 adjacent normal breast tissue (missing two cases of invasive carcinoma, so we have 108 cases in total for analysis). After deparaffinization and rehydration, antigen retrieval was done by Tris-EDTA Buffer (pH 9.0). The samples were treated with methanol with 1% hydrogen peroxide for 30 min to block endogenous peroxidase activity, followed by incubation with 10% normal goat serum to prevent nonspecific staining; then, the samples were incubated with anti-PARG antibody (1:50. Cell signaling Technology, 24489S) at 4 °C overnight. The samples were incubated with a biotinylated secondary antibody (1:200. Vector Laboratories, PK-6101) for 30 min at ambient temperature and with avidin–biotin peroxidase complex solution (1:100) for additional 30 min at ambient temperature. The slides were incubated with DAB solution and subsequently counterstained with haematoxylin. The slides were scanned on the Vectra Polaris Automated Quantitative Pathology Imaging System (PerkinElmer, Waltham,US) for quantification by Visiopharm platform (Visiopharm, Hoersholm, Denmark). A total score of protein expression was calculated from both the percentage of immunopositive cells and the immunostaining intensity. High and low protein expressions were defined using the mean score of all normal tissues as a cutoff point.

DePARylation activity of PARG in whole cell lysates

The dePARylation activity of PARG was measured in the PARP1 Histone H4 Activity Assay (#K611, Tulip BioLabs). Briefly, wells of the plate were incubated with activated PARP1 mixtures to generate PARylated Histone H4 or PARP1 mixtures without NAD+, which were set as the Blank, according to the manufacturer instruction. After that, the plate was washed with PBS three times. Equal number of cells prepared from different cell lines were lysed by equal amount of NETN buffer in the presence of 10 µM olaparib, to inhibit the endogenous PARP1 activity, on ice for 30 min. After centrifugation at 14,300xg at 4 °C for 10 min, supernatant was kept, and the protein concentration was determined by BCA assay. Wells of the plate were either incubated with equal amount of whole cell lysates in NETN buffer with 10 µM olaparib or just NETN buffer with 10 µM olaparib, set as the Control, overnight at room temperature. The Blank was also incubated with equal NETN buffer with 10 µM olaparib. After that, the plate was washed with PBS three times. The remnant pADPr were detected by anti-pADPr antibodies and quantified by TMB substrate with the OD450 read, according to the manufacturer instruction. The percentage of pADPr hydrolysis was calculated by normalization with the Blank and the Control.

Removed in the revised version: Results

Removed in the revised version: PARG depletion leads to drastic sensitivity to PARGi

Removed in the revised version: The results of early studies indicated that defective DNA repair pathways and DNA replication stress would result in enhanced sensitivity to PARG inhibitor, PDD00017273 (PARGi)25–27, 38, 39. To further define the key DDR pathways that are important for cellular response to PARGi, we performed CRISPR screening using our homemade DDR sgRNA library, which targets approximately 360 genes involved in various DDR pathways34. Interestingly, our screen results showed that PARG depletion led to significantly increased cellular sensitivity to PARGi (Fig. 1A, Table S1). In addition, we found that POLB (Fig. 1A), which is important for BER, a repair pathway that relies on PARP1 functions and pADPr-dependent DNA damage signaling, also showed synthetic lethality with PARGi. The synthetic lethality between POLB and PARGi was consistent with the findings of a previous report38, indicating the high quality of our screening results.

Removed in the revised version: PARG loss sensitizes cells to PARGi.

Removed in the revised version: a, Ranking of co-essential genes with PARGi treatment on the basis of a DrugZ analysis of the results of CRISPR/Cas9 screening with a DDR library in HEK293A cells. The NormZ score was used to determine a possible synthetic lethality gene under PARGi treatment. Drug-sensitive genes were marked in red; drug-resistant genes were marked in blue on the basis of the false discovery rate (FDR, 0.05 cut-off). b, HEK293A WT cells and HEK293A PARG KO cells were treated with different doses of PARGi for 72 hours. Cell viability was determined by the CellTiter-Glo assay. c, HEK293A PARG KO cells, re-constituted with either full-length PARG or catalytic domain mutation of PARG, were treated with different doses of PARGi for 72 hours. Cell viability was determined by the CellTiter-Glo assay.

Removed in the revised version: To further validate our data, we generated PARG KO cells in the same 293A cells used in DDR sgRNA library screening. Interestingly, PARG KO cells showed extreme sensitivity to PARGi, with nearly thousand-fold sensitivity (the IC50 in 293A cells was 96±24 µM and 210±30 nM in PARG KO cells, on the basis of five independent biological replicates) (Fig. 1B). In addition, we used another two structurally similar compounds (PDD00017272 and PDD00017238) to further confirm the extreme sensitivity of PARG KO cells to PARG inhibition (Fig. S1B), while the inactive but structurally similar compound PDD0031705 did not distinguish wild-type (WT) or PARG KO cells (Fig. S1C). To further validate that this sensitivity was not restricted to one cell line, we generated HeLa-derived PARG KO cells with a different gRNA. As expected, HeLa-derived PARG KO cells also displayed extreme sensitivity to PARGi (Fig. S1D). These results suggest that PARG depletion resulted in drastic sensitivity to PARGi.

Removed in the revised version: We reconstituted PARG KO cells with WT PARG or a PARG catalytic-inactive mutant E755/756A40 to further investigate the nature of PARGi sensitivity in PARG KO cells. Cells treated with the alkylating agent methyl methanesulfonate (MMS), which is known to activate pADPr signaling, were used to test PARG activity. As shown in Fig. S1E, PARG KO cells and cells reconstituted with PARG catalytically inactive mutant showed a significant increase in pADPr signaling induced by MMS, while WT cells or cells reconstituted with WT PARG did not. Expectedly, reconstitution with WT PARG reversed the sensitivity of PARG KO cells to PARGi, while reconstitution with PARG catalytic inactive mutant failed to do so (Fig. 1C). Together, these data strongly suggest that loss of PARG activity induces dramatic sensitivity to PARGi.

Removed in the revised version: The question is why PARG KO cells would be sensitive to PARGi. There are at least two potential explanations for this observation. First, PARGi may target not only PARG but also a second target, which only becomes essential in the absence of PARG. Second, our PARG KO cells, despite being generated with two different sgRNAs and in two independent cell lines validated by Western blotting and DNA sequencing, are not complete PARG KO cells. The two different sgRNAs used here target all three catalytically active isoforms (isoforms 1, 2 and 3), and sgRNA#2 used in HeLa cells also targets isoforms 4 and 5, but these two isoforms are considered catalytically inactive according to the Uniprot database (Fig. S1D). To address this important question, we performed several experiments to test both hypotheses. Please see below for additional studies.

Removed in the revised version: PARGi induces PARP1/2-dependent cell death in PARG KO cells

Removed in the revised version: As shown above (Fig. S1E), MMS treatment significantly increased pADPr signaling in PARG KO cells. Interestingly, pADPr signaling further increased modestly in PARG KO cells treated with the combination of PARGi and MMS (Fig. S2A). To further evaluate the underlying mechanisms of PARGi sensitivity in PARG KO cells, we generated PARP1/2 DKO cells in both WT and PARG KO cells and treated them with PARGi (Fig. 2A). PARP1 cleavage was detected in PARG KO cells treated with PARGi, indicating apoptosis. Moreover, pADPr accumulated significantly in PARG KO cells treated with PARGi, but not in other cells (Fig. S2B). Furthermore, we used flow cytometry to detect pADPr signaling following short treatment with PARGi (Fig. 2B). Consistently, moderate pADPr signaling was observed in PARG KO cells, but not in other cells (Fig. 2B).

Removed in the revised version: PARGi treatment induces NAD+- and PARP-dependent cell death in PARG KO cells.

Removed in the revised version: a, HEK293A WT, PARG KO, PARP1/2 DKO, and PARG/PARP1/2 TKO cells were treated with PARGi (1µM) for 72 hours. The total cell lysates were immunoblotted with the indicated antibodies. b, HEK293A WT, PARG KO, PARP1/2 DKO, and PARG/PARP1/2 TKO cells were treated with DMSO or 10 µM PARGi for 4 hours and then fixed and stained with anti-pADPr antibody and propidium iodide (PI). c, HEK293A WT, PARG KO, PARP1/2 DKO clls, and PARG/PARP1/2 TKO cells were treated with different doses of PARGi for 72 hours. Cell viability was determined by the CellTiter-Glo assay. d, Relative NAD+ level in HEK293A WT and PARG KO cells with the indicated treatment for 48 hours. PARGi, 10 µM; PARPi, 10 µM; NAM, 100 µM; NMN, 1 mM; FK866, 10 nM. e, PARG KO cells were treated with PARGi (10 µM) or PARGi and NAM (100 µM) or NMN (1mM) for 48 hours. The total cell lysates were immunoblotted with the indicated antibodies. f, Results of clonogenic assays conducted using HEK293A PARG KO cells treated with PARGi (500 nM) or PARGi and NAM (100 µM) or NMN (1 mM) for 7 days.

Removed in the revised version: We reason that sensitivity to PARGi in PARG KO cells may depend on PARP1/2-dependent pADPr accumulation. To test this hypothesis, CellTiter-Glo and a colony formation assay were used to detect the sensitivity of WT, PARG KO, PARP1/2 DKO, and PARG/PARP1/2 TKO cells to PARGi. As expected, the sensitivity of PARG KO to PARGi was completely reversed in TKO cells, and DKO cells showed moderate resistance compared with WT cells (Fig. 2C). Indeed, TKO cells completely rescued cell death induced by PARGi in PARG KO cells (Fig. S2C). In addition, treatment with PARPi olaparib was able to reverse PARGi sensitivity in PARG KO cells (Fig. S2C). These data together suggest that the cytotoxicity of PARGi in PARG KO cells was due to uncontrolled PARylation, which requires PARP1/2.

Removed in the revised version: PARP1/2 catalyzes the pADPr reaction by transferring the ADR-ribose moiety of NAD+ to the acceptor proteins with the release of NAM41. Uncontrolled PARylation may lead to cytotoxicity, at least in part because of NAD+ depletion23, 42, 43. To determine whether cell death induced by PARGi in PARG KO cells is at least partially caused by NAD+ depletion, we measured the relative NAD+ level in WT and PARG KO cells treated with PARGi and used the nicotinamide phosphoribosyltransferase (NAMPT) inhibitor FK866 as a positive control. FK866 is a noncompetitive inhibition of NAMPT, a key enzyme involved in the regulation of NAD+ biosynthesis from its natural precursor NAM 44. As expected, the NAD+ level decreased modestly over the treatment period in PARG KO cells, but this decrease was not as dramatic as those observed in cells treated with FK866 (Fig. S2D). Moreover, treatment with FK866, which significantly decreased the NAD+ level, did not lead to cell death (Fig. S2E), suggesting that the decreased NAD+ level is not sufficient or the only reason for the cytotoxicity observed in PARG KO cells treated with PARGi. Furthermore, we treated cells with the NAD+ precursors nicotinamide mononucleotide (NMN) or NAM23, 45. As expected, NMN, NAM, and PARPi rescued the NAD+ decrease in both non-treated and PARGi treated cells (Fig. 2D). In addition, NMN and NAM were able to not only reduce pADPr accumulation induced by long-term PARGi treatment (Fig. 2E), but also rescue cell lethality caused by PARGi treatment in both HEK293A-derived and HeLa-derived PARG KO cells (Fig. 2F and S2F). Together, these results suggest that uncontrolled PARylation in the absence of any dePARylation activities would lead to cytotoxicity, probably due to several mechanisms including apoptosis and are not limited to NAD+ depletion/reduction.

Removed in the revised version: Unexpected S phase pADPr signaling observed in PARG KO cells treated with PARGi

Removed in the revised version: PARP1/2 are activated by DNA damage to generate pADPr signals. Under normal growth conditions, PARP1/2 activities may be low; therefore, it was extremely difficult to detect pADPr. Unexpectedly, pADPr signaling from normal proliferating cells without any exogenous DNA damage was detected in PARG KO cells after PARGi treatment (Fig. 2B and S2B), which was apparent in S phase cells, as detected by FACS analysis (Fig. 2B). To further characterize pADPr and DNA damage signaling induced by PARGi or other DNA damaging agents, we used flow cytometry to detect both pADPr and γH2AX signaling (Fig. 3A and S3A). Interestingly, pADPr signaling was specifically detected in S phase in PARG KO cells treated with PARGi, while at the same time, DNA damage signaling only showed a slight increase or was not observed, as revealed by anti-γH2AX antibody (Fig. 3A). On the other hand, MMS treatment led to increased pADPr signaling throughout the cell cycle (Fig. 3A). However, it only led to an S phase-specific increase in DNA damage signaling, i.e. γH2AX. Our interpretation of the MMS results is that MMS treatment would lead to DNA alkylation throughout the cell cycle, which is repaired by the BER pathway. BER creates single-strand nicks as repair intermediates, which are recognized by PARP1 and activate PARP1. Therefore, pADPr signaling can be detected in all cell cycle phases in PARG KO cells under this condition. However, single-strand nicks are not recognized by DNA damage checkpoint pathways. These nicks, when encountering replication forks, will be converted to DSBs and therefore activate DNA damage checkpoint kinases, resulting in increased γH2AX signaling specifically in S phase cells.

Removed in the revised version: PARGi treatment induces S phase-specific pADPr signaling in PARG KO cells.

Removed in the revised version: a, HEK293A WT and PARG KO cells were treated with DMSO or 10 µM PARGi for 4 hours or 0.01% MMS for 30 min and then fixed and stained with anti-pADPr antibody or anti-γH2A.X antibody and PI. b, HEK293A WT and HEK293A PARG KO cells were mock-treated or pre-treated with 2 µM emetine for 90 min and then treated with PARGi for an additional 4 hours. Cells were fixed and stained with anti-pADPr antibody and PI. c, HEK293A PARG KO cells were synchronized with double thymidine block (DTB). Cells remained with DTB or were released from DTB, treated with 10 µM PARGi for 4 hours, and then fixed and stained with anti-pADPr antibody and PI. d, Representative images and results (left) of clonogenic assays conducted using control cells and DTB synchronized or released HEK293A PARG KO cells treated with the indicated doses of PARGi for 7 days, and quantification of crystal violet staining assay (right). e, Results of clonogenic assays were conducted in PARG KO cells with indicated treatment for 7 days (PARGi, 1 µM; adarotene, 200 nM; CD437, 800 nM).

Removed in the revised version: To further explore the relationship between pADPr and γH2AX signaling, we conducted a similar flow cytometry analysis in parental WT and PARG KO cells following treatments with different DNA damaging agents. Again, PARGi treatment in PARG KO cells led to an S phase-specific increase in pADPr signaling; this unique pattern was not detected in response to any of the DNA damaging agents used in this study (Fig. S3A). Our current working hypothesis is that PARG KO cells treated with PARGi may completely block any dePARylation reactions. The dramatic S phase-specific pADPr signaling detected under this condition indicates that PARP1 is specifically activated in S phase cells.

Removed in the revised version: During DNA replication, maturation of Okazaki fragments involves single-strand nicks on DNA, which is normally ligated by LIG1. These unligated Okazaki fragment intermediates could activate PARP1 to generate endogenous S phase pADPr6. It is possible that in such situations, the maturation of Okazaki fragments may require PARP1-dependent recruitment of XRCC1/LIG36, which agrees with the results of early studies that indicate that both LIG1 and LIG3 are required for Okazaki fragment maturation46. Of course, PARP1 may have other functions in the S phase, including DNA repair47.

Removed in the revised version: The S phase pADPr signaling observed in PARG KO cells after PARGi treatment was observed under conditions in which no or very limited DNA damage or replication stress– induced signaling was detected by anti-γH2AX antibody (Fig. 3A) and S phase pADPr resulting from unligated Okazaki fragments were observed under similar conditions6; thus, we reasoned that unligated Okazaki fragments in normal proliferating cells may render this unexpected S phase pADPr signaling in PARG KO cells after PARGi treatment. Indeed, pre-treatment with emetine, which diminishes Okazaki fragments by its anti-protein synthetic activity6, 48, 49, greatly inhibited S phase pADPr signaling in PARG KO cells (Fig. 3B). Furthermore, mild inhibition of POLA1 with two inhibitors (i.e. adarotene and CD437) also abolished the PARGi induced S phase pADPr signaling in PARG KO cells (Fig. S3B).

Removed in the revised version: To further support S phase-specific pADPr signaling resulting from unligated Okazaki fragments, we performed double thymidine block (DTB) and release experiments. Indeed, release into S phase is critical for the pADPr signaling observed in PARG KO cells treated with PARGi (Fig. 3C). Moreover, the cytotoxicity of PARGi also requires S phase progression, since both control proliferating cells and double thymidine blocked and released cells died following PARGi treatment, while cells arrested by DTB failed to do so (Fig. 3D). Additionally, mild inhibition of POLA1 partially rescued the cytotoxicity of PARGi in PARG KO cells (Fig. 3E).

Removed in the revised version: PARG KO cells show prolonged PARP1 chromatin binding

Removed in the revised version: Following DNA damage, PARP1 detects and is activated by both SSBs and DSBs to initiate subsequent DNA repair50, 51. After that, PARP1 dissociates from DNA. If this does not occur, the persistent PARP1-DNA complexes could trap PARP1 on chromatin, which is one of the key characteristics of PARPi-mediated cytotoxicity35. Given that the sensitivity of PARG KO cells to PARGi resulted from PARP1/2-dependent pADPr accumulation, which indicates an association and activation of PARP1/2 by DNA, we used a previously described trapping assay to measure the levels of chromatin-bound PARP135, 36. Unexpectedly, the levels of chromatin-bound PARP1/2 increased in PARG KO cells, regardless of treatment, while the soluble fraction of PARP1 decreased compared with that in WT cells (Fig. S4A).

Removed in the revised version: To further explore the chromatin-bound PARP1, we performed similar soluble and chromatin fractionation and enrichment of PARylated proteins by Af1521 beads. Consistently, chromatin-bound PARP1 increased in PARG KO cells, regardless of whether these cells were treated with PARGi (Fig. 4A). Moreover, PARylated PARP1 also increased in chromatin fraction. Interestingly, following PARGi treatment, PARG KO cells showed more pADPr levels than did WT cells in chromatin fraction, while similar pADPr levels were detected in soluble fractions (Fig. 4A). Indeed, a recent study by Gogola and colleagues suggest that PARG depletion does not enhance PARP1 dissociation from DNA but prevents excessive PARP1 binding, on the basis of the results of a similar trapping assay and a laser-induced DNA damage assay for measuring PARP1 association52. Taken together, these results indicate that prolonged PARP1 chromatin binding or trapping occurs in PARG KO cells, which may contribute to the sensitivity of these cells to PARGi.

Removed in the revised version: Prolonged PARGi treatment induces pADPr throughout the cell cycle and DDR in PARG KO cells.

Removed in the revised version: a, Immunoblots of chromatin-bound PARP1 and PARylated proteins in HEK293A WT and PARG KO cells treated with PARGi (10 µM) for 4 hours. PARylated proteins were enriched by Af1521 beads. b, Sensitivity of HEK 293A PARG/PARP1/2 TKO and PARP1 reconstitution cells to PARGi. Cells were treated with different doses of PARGi for 72 hours, and cell viability was determined by the CellTiter-Glo assay. c, Prolonged PARGi treatment induces pADPr throughout the cell cycle in PARG KO cells. HEK293A WT and PARG KO cells were treated with PARGi (10 µM) for the indicated time and then fixed and stained with anti-pADPr antibody and violet. d, Immunoblotting of γH2A.X signals and other indicated proteins and modifications induced by prolonged PARGi+/-MMS treatment.

Removed in the revised version: To further support the contribution of chromatin-bound PARP1 to PARGi sensitivity, we introduced either WT PARP1 or a trapping-deficient PARP1del.p.119K120S construction into TKO cells (Fig. S4B). PARP1del.p.119K120S mutant is a trapping-deficient mutant, which leads to PARPi resistance53, 54. Notably, expression of WT PARP1 led to dramatic PARGi sensitivity in TKO cells, while PARP1del.p.119K120S mutant only showed a slight increase in PARGi sensitivity (Fig. 4B). Together, these data suggest that increased chromatin-bound and trapped PARP1 contributes, at least in part, to the sensitivity of these cells to PARGi.

Removed in the revised version: To investigate the consequence of S phase pADPr and prolonged PARP1 chromatin binding in PARG KO cells, we again used flow cytometry to detect pADPr signaling in a time course following PARGi treatment. Consistent with the aforementioned data, S phase pADPr was detected following short PARGi treatment (i.e. 4 hours), while pADPr signaling was detected throughout the cell cycle following prolonged PARGi treatment, i.e. 24 or 48 hours (Fig. 4C). Furthermore, SSBs were not detected following short PARGi treatment, as measured by an alkaline comet assay (Fig. S4C), which may due to the sensitivity of this assay. However, prolonged PARGi treatment in PARG KO cells led to dramatic increase of SSBs detected by this assay (Fig. S4C). Similarly, we observed S phase-specific pADPr signaling with short PARGi treatment and pADPr signaling throughout the cell cycle with prolonged PARGi treatment in HeLa PARG KO cells, although significant S-phase pADPr was detected in control WT HeLa cells (Fig. S4D), which was not observed in control WT 293A cells (Fig. 4C). The difference in pADPr signaling between HeLa and 293A cells may due to a low level of PARG and/or high activity of PARP1 in HeLa cells, in which high–molecular weight smears were detected following PARGi and MMS treatment (Fig. S4E). Together, these data indicate that PARG regulates S-phase pADPr in normal proliferating cells.

Removed in the revised version: To further confirm pADPr signaling observed by flow cytometry analysis, we detected pADPr and γH2AX by Western blotting in HEK293A and HeLa cells and their corresponding PARG KO cells, while additional MMS treatment was included as a positive control (Fig. 4D and S4F). Expectedly, we only observed increased γH2AX signaling in PARG KO cells with prolonged PARGi treatment (48 hours of treatment in 293A PARG KO cells and 24 and 48 hours of treatment in HeLa PARG KO cells) or those treated with additional MMS, indicating that pADPr accumulation leads to DNA damage throughout the cell cycle. The pADPr signaling in PARG KO cells pre-treated with PARGi was further enhanced in MMS-treated cells (Fig. 4D and S4F), likely due to increased SSBs; nevertheless, PARGi-mediated pADPr accumulation in PARG KO cells under condition without any exogenous DNA damage was sufficient to cause cytotoxicity. This situation is different from the hyperactivity and progressive inactivation of PARP1 in XRCC1-deficient cells after MMS treatment42. Taken together, we speculate that inhibition of dePARylation in PARG KO cells treated with PARGi lead to pADPr accumulation and cell lethality.

Removed in the revised version: CRISPR screens reveal genes responsible for regulating pADPr signaling and/or cell lethality in WT and PARG KO cells

Removed in the revised version: PARG KO and PARGi appear to have a synergic effect on cell viability, which depends on PARP1/2-dependent pADPr accumulation (Fig. 1 and 2). Further, we showed that unligated Okazaki fragments were the likely source of SSBs, which leads to S phase pADPr signaling (Fig. 3). Moreover, we found that uncontrolled S phase pADPr accumulation eventually leads to DNA damage, as revealed by anti-γH2AX antibody and alkaline comet assay (Fig. 4). These data prompted us to further explore potential genes that may be responsible for controlling cell viability and pADPr signaling in control WT and PARG KO cells, with or without PARGi treatment. As we speculated above, there may be a potential second gene targeted by PARGi, which only becomes essential in the absence of PARG. We thus performed the outlined whole-genome CRISPR screens to address these questions (Fig. 5A), similar to the studies published by us and others37, 55.

Removed in the revised version: CRISPR screening identifies regulators of pADPr and cell viability.

Removed in the revised version: a, Workflow of whole-genome CRISPR screens. For FACS-based CRISPR screening, 5 days after puromycin selection, cells were treated with PARGi (10uM) for 4 hours and then stained with anti-pADPr antibody and sorted with flow cytometry. Cells with strong signals (top 25%, TOP) and weak signals (bottom 25%, BOT) were selected. The sgRNAs from these cells were then sequenced and analyzed. For cell viability screening, cells were treated with or without PARGi for 21 days before collection. b, Scatter plot of DrugZ scores of PARG KO cells treated with or without PARGi treatment. The genes in the same pathway were marked with specific colors. A positive score indicates an enhanced pADPr signal, while a minus score indicates a decreased pADPr signal. c and d, Ranking of PARGi co-essential genes on the basis of a DrugZ analysis of the results of CRISPR/Cas9 screens performed with Toronto Knock Out Library (version 3) in HEK293A PARG KO cells and HEK239A cells. e, Analysis of biological processes of PARGi co-essential genes identified in HEK293A PARG KO cells and HEK239A cells. f, HEK293A WT cells, PARG KO, POLB KO, LIG1 KO, XRCC1 KO, and LIG3 knockdown cells were treated with different doses of PARGi for 72 hours. Cell viability was determined by the CellTiter-Glo assay. g, The cell Viability of HEK293A WT and PARG cells under POLB knockdown to PARGi. Cells were treated with different doses of PARGi for 72 hours.

Removed in the revised version: As for FACS-based CRISPR screening, we showed that PARP1 consistently appeared as the gene that, when depleted, led to reduced pADPr in all settings (Fig. 5B, S5A and Table S2-4), which agrees with the dominant role of PARP1 in promoting PARylation in cells. The depleted genes that led to increased pADPr signaling were also similar in all of these settings (Fig. 5B and S5A); these include BER genes (i.e. POLB, XRCC1, LIG3, and CHD1L) and genes involved in DNA replication/Okazaki fragment maturation (i.e. RFCs, FEN1, and LIG1). Interestingly, ARH3 only appeared in PARG KO cells treated with PARGi (Fig. 5B and S5A), indicating that as anticipated ARH3 may serve as a backup dePARylating enzyme.

Removed in the revised version: We further investigated whether the aforementioned genes that regulated pADPr signaling would show strong synthetic lethality with PARGi in both PARG KO cells (Fig. 5C and Table S5) and WT cells (Fig. 5D and Table S6). To identify the most confident hits, we used a stringent false discovery rate (FDR) threshold of 0.0556. Consistently, PARP1 was listed as the most resistant gene with PARGi in PARG KO cells (Fig. 5C). Interestingly, ARH3 appeared as the top gene in PARG KO cells treated with PARGi (Fig. 5C), while other dePARylation enzymes (e.g. MARCROD1/2 andTARG1/OARD1) did not show any synthetic lethal effect in these screens. These data agree with the results of FACS-based screens and suggest that ARH3 is a backup enzyme in the absence of PARG following PARGi treatment. Indeed, loss of ARH3 further sensitized PARG KO cells to PARGi (Fig. S5B). In addition, HPF1, another gene/protein involved in PARylation regulation and licensing serine mono-ADP-ribosylation by PARP1/257, was among the top synthetic lethality genes in PARG KO cells treated with PARGi, but not in WT cells (Fig. 5C,D and S5C).

Removed in the revised version: Thus, our CRISPR screens revealed that genes involved in PARylation regulation (ARH3 and HPF1), BER (i.e. POLB, XRCC1, LIG3, and CHD1L), and DNA replication/Okazaki fragment maturation (i.e. RFCs, FEN1, and LIG1) are responsible for pADPr signaling and/or cell lethality in WT and PARG KO cells.

Removed in the revised version: To further explore the involvement of these aforementioned genes, we compared synthetic lethality genes in PARG KO cells with those in WT cells treated with PARGi. Interestingly, two BER key genes, POLB and XRCC1, showed strong synthetic lethality with PARGi in both PARG KO cells (Fig. 5C) and WT cells (Fig. 5D), which is consistent with the results of a previous report that POLB-deficient cells are sensitive to PARGi38. Moreover, the GO analysis revealed that the BER pathway displays synthetic lethality in both PARG KO cells and WT cells following PARGi treatment (Fig. 5E). Importantly, LIG1 and LIG3, the two ligases that are directly involved in Okazaki fragment ligation, were also listed as synthetic lethality genes (Fig. 5C and 5D). To further validate these data, we created several KO/KD cell lines, including KO of BER genes (i.e. POLB and XRCC1) and genes involved in Okazaki fragment ligation (i.e. LIG1 and LIG3) (Fig. S5D). Consistent with previous reports, XRCC1 loss destabilizes LIG3 (Fig. S5D)58, 59. Notably, loss of genes involved in Okazaki fragment ligation induced S phase pADPr signaling, just like PARG KO cells following PARGi treatment, while loss of BER genes resulted in pADPr signaling throughout the cell cycle (Fig. S5E). These results agree with the aforementioned results that unligated Okazaki fragments are the likely source of SSBs, which induces S phase-specific pADPr signaling. Also consistent with CRISPR screening data, loss of genes involved in Okazaki fragment ligation or BER resulted in increased sensitivity to PARGi (Fig. 5E). Furthermore, we knocked down POLB in WT and PARG KO cells (Fig. S5F), which resulted in increased sensitivity to PARGi in both WT and PARG KO cells (Fig. 5G). Together, our data indicate that unligated Okazaki fragments induce S phase-specific pADPr, while BER intermediates lead to pADPr throughout the cell cycle; deficiency in either of these events would increase cytotoxicity following PARGi treatment.

Removed in the revised version: PARG expression is a potential biomarker for PARGi-induced cytotoxicity

Removed in the revised version: We showed that cells with no detectable full-length PARG expression are extremely sensitive to PARG inhibition (Fig. 1 and S1). Thus, we are interested in determining whether PARG expression would correlate with PARGi sensitivity. An early study examined a panel of ovarian cancer cell lines and reported PARGi sensitivity in a subset of these cell lines39. We compared PARG expression in PARGi sensitive and resistant cell lines and found that PARG expression was significantly reduced in PARGi sensitive cell lines (Fig. 6A). To support this result, we knocked down PARG by shRNA in HeLa cells and showed that loss or further reduction of PARG led to dramatic sensitivity to PARGi (Fig. 6B). To further test our hypothesis, we ranked the PARG mRNA level in several ovarian cancer cell lines using the CCLE database (Fig. 6C). We identified an ovarian cancer cell line, RMUGS, which has low PARG expression, similar to that observed in another known PARGi sensitive cell line, KURAMOCHI. Notably, RMUGS cells were considered as potential PARGi-resistant cells based on an earlier report39. However, we showed that RMUGS cells were very sensitive to PARGi treatment (Fig. 6D), and the sensitivity of these cells to PARGi could be reversed by PARPi treatment (Fig. 6E). These data indicate that PARG expression may dictate the sensitivity of tumor cells to PARGi-based therapy.

Removed in the revised version: PARG expression is a potential marker for PARGi sensitivity.

Removed in the revised version: a, PARG mRNA level comparison between sensitive (RMGI, KURAMOCHI, and OVMANA) and resistant cells (COV362, COV318, OV56, OVISE, OVSAHO, CAOV3, and OVCAR3) from Pillay’s work on the basis of CCLE data. b, Clonogenic assay results of control and PARG knockdown HeLa cells with PARGi (2 µM) treatment for 7 days. shRNA knockdown efficiency was confirmed by an immunoblot of PARG. c, Ranked PARG expression level in ovarian cancer cell lines based on the CCLE database. The sensitive and resistant cells from Pillay’s work were labeled. d, Clonogenic assay results of OVCAR3 and RMUGS treated with or without PARGi (2 µM). e, Clonogenic assay results of RMUGS treated with PARGi (2 µM), PARPi (2 µM), or both. f, Left: Representative images of PARG IHC staining in breast and ovarian tissues and tumor samples to determine the PARG expression level. The summary is listed at the bottom. Scale bar, 200 µm. Right: The scatter plot of mean immunostaining intensity of PARG in each sample. The mean of each group was plotted.

Removed in the revised version: An early study reported that loss of PARG expression resulted in PARPi resistance52. They demonstrated that PARG is frequently lost in acquired PARPi-resistant mouse mammary tumors and further revealed that PARG depletion occurs in triple-negative breast and ovarian cancer. Another study showed that 60 of 274 (22%) of human ovarian tumors have low PARG expression38. Thus, PARGi can be potentially used to treat these PARPi-resistant cancers, especially those with low PARG expression. We further performed an IHC assay to detect PARG expression in breast and ovarian cancer TMA samples. The IHC assay was established with the use of our xenograft tumors derived from HeLa and HeLa PARG KO cells (data not shown). As shown in Figure 6F, PARG expression was detected in normal ovarian and breast tissues. PARG expression was also detected in a majority or most of ovarian and breast tumor samples (Fig. 6F). Interestingly, about 5% to 40% of these breast or ovarian cancer tumor samples showed low expression levels of PARG, some even showed no detectable level of PARG, compared with those in normal tissue samples (Fig. 6F). These data suggest that PARG downregulation occurs in breast and ovarian cancers. PARGi may be a promising strategy for the treatment of these cancers with low PARG expression.

Removed in the revised version: In early studies, HR-deficient cells showed increased sensitivity to PARG inhibition18, 25, 26, 60. Similarly, we observed PARGi-induced cytotoxicity in RPE1 Flag-Cas9 TP53/BRCA1 DKO cells, but not in control cells (Fig. S6A). Moreover, combining PARPi and PARGi did not reveal any additive effect (Fig. S6A). Furthermore, we created inducible BRCA1 depletion cells using the auxin-inducible degron (mAID) tag. As shown in Figure S6B, HR-proficient cells showed normal sensitivity to PARGi, while HR deficiency by the loss of BRCA1 conferred modest sensitivity to PARGi. The PARGi sensitivity in BRCA1 depletion cells was reversed to normal by the restoration of HR due to 53BP1 loss. Consistent with the results obtained in RPE1 cells, the PARPi+PARGi combination did not lead to any further change in cytotoxicity. While these results indicate that PARGi can be used to target HR-deficient cancers, the effect is quite modest. Moreover, PARGi treatment did not show any effect on PARPi-resistant cells due to restoration of HR mediated by 53BP1 loss.

Removed in the revised version: To further investigate the potential contribution of HR deficiency and/or PARG loss to PARGi sensitivity, we knocked down PARG by shRNA in these inducible BRCA1 depletion cells with or without 53BP1 KO (Fig. S7A). As expected, PARG loss is the main driver of PARGi sensitivity; regardless whether these cells are HR-proficient, HR-deficient, or HR-restored because of 53BP1 loss (Fig. S7B). Our data therefore strongly suggest that PARG expression is the major determinant of PARGi sensitivity.

Removed in the revised version: PARG is essential for cell survival

Removed in the revised version: Our results suggested that cells with no detectable full-length PARG expression or low expression of PARG are extremely sensitive to PARG inhibition (Fig. 1 and 6). As aforementioned, the question is why PARG KO cells would be sensitive to PARGi. To address this question, we first performed whole-genome CRISPR screening. However, the potential second target was not revealed by our CRISPR screening data (Fig. 5). The second possibility is that our PARG KO cells are not complete PARG KO cells. These cells may have residual PARG expression or activity and only cells with very low PARG expression are sensitive to PARGi, as shown in Fig. 6. Indeed, two independent Parg knockout mice have been created by targeting early exons of PARG13, 61. While one group of Parg knockout mice showed embryonic lethality and cells derived from these mice only survived in the presence of PARPi13, the other Parg knockout mice did not display any major phenotypes, probably because truncated Parg and Parg activity were still detected61. We favor the second possibility, especially since the PARG catalytic domain is at its C-terminus.

Removed in the revised version: To test this possibility, we first used two independent antibodies that recognize the C-terminus of PARG to detect potential residual PARG isoforms or spliced variants (Fig. S9A). Unfortunately, besides full-length PARG, these antibodies also recognized several other bands, some of them were reduced or absent in PARG KO cells, others were not. Thus, we could not draw a clear conclusion which functional isoform/truncated form was expressed in our PARG KO cells. Then, we used two additional gRNAs which target respectively the beginning of the catalytic domain and the sequence around the catalytic site (Fig. 7A). Interestingly, we were only able to obtain viable clones with no detectable PARG in the presence of PARPi in either 293A or HeLa cells, which we named as PARG complete/conditional KO cells (cKO) (Fig. 7B). Indeed, these cKO cells could not survive without PARPi (Fig. 7C). To further investigate whether cell lethality is associated with PARG activity in these cells, we reconstituted PARG cKO cells with WT PARG or a PARG catalytic-inactive mutant. Consistent with previous result (Fig. 1 and S1), WT PARG not only reduced pADPr level but also rescued cell lethality due to PARG loss, while the catalytic-inactive mutant PARG failed to do so (Fig. 7D).

Removed in the revised version: PARG is essential for cell survival.

Removed in the revised version: a, Diagram of full-length PARG was presented with the indicated gRNAs (gRNA#3 and 4) which target different regions in the C-terminal catalytic domain. The boundary of the catalytic domain was depicted based on Uniprot annotation. gRNA#1 and gRNA#2 were used previously to generate the aforementioned HEK293A- and HeLa-derived PARG KO cells, respectively; while gRNA#3 and gRNA#4 were used to generate PARG complete/conditional knockout (cKO) in the presence of olaparib in HEK293A and HeLa cells. b, Immunoblotting was conducted to confirm the loss of PARG in PARG cKO cells derived from HEK293A and HeLa cells, which were cultured in the presence of 100 nM olaparib. c, Clonogenic assay results of WT and PARG cKO cells treated with or without PARPi (100nM) for 7 days. d, Left: The immunoblots to confirm reconstitution with WT PARG or catalytic inactivation PARG in HEK293A PARG cKO cells. Right: Results of clonogenic survival assay with HEK293A PARG cKO cells reconstituted with WT or catalytic inactivation mutant of PARG for 7 days. e, Representative clonogenic results conducted in HEK293A PARG cKO cells treated with NAM (100 µM) or NMN (1 mM) for 7 days. f, HEK293A PARG cKO cells were synchronized with double thymidine block (DTB). Cells remained with DTB or were released from DTB for 4 hours, and then fixed and stained with anti-pADPr antibody and FxCycle Violet dye. g, Immunoblots of soluble and chromatin-bound PARP1 and pADPr levels in HEK293A WT and PARG cKO cells treated with DMSO or olaparib (10 µM) for 2 hours.

Removed in the revised version: The data above indicate that our PARG KO cells have residual PARG activity and treatment with PARGi in these cells further decreases PARG activity, which mimic the complete loss of PARG as that in PARG cKO cells. To further test this hypothesis, we measured the relative PARG dePARylation activity in whole cell lysates (WCL) prepared from different cell lines (Fig. S9C). Consistent with our working hypothesis, residual PARG activity was detected in WCL prepared from PARG KO cells, while PARG activity was barely detected or absent in WCL isolated from PARG cKO cells. Please note that we incubated cell lysates with substrates overnight to evaluate the maximum level of pADPr hydrolysis, i.e. PARG activity, we were able to detect in these assays. It is very likely that the PARG activity in PARG KO cells was much lower than that indicated in Fig. S9C, due to saturation of signals for lysates isolated from wild-type cells. Thus, the data presented here may under-estimate the reduction of PARG activity in PARG KO cells. Nevertheless, these data indicate the residual PARG activity in PARG KO cells, which is absent in PARG cKO cells. Furthermore, the PARG activity was further inhibited by PARGi in a dose-dependent manner in WCLs prepared from both WT and PARG KO cells (Fig. S9D-E). Notably, the PARG activity was consistently lower in WCL prepared from PARG KO cells when compared with that in control wild-type cells in the presence of a wide range of PARGi concentrations. However, a significant fraction of PARG activity still existed in WCL prepared from wild-type cells even at the highest concentration of PARGi used in this study (Fig. S9E). These data indicate residual PARG activity in PARG KO cells likely accounts for the survival of these cells. Moreover, further inhibition of this activity by PARGi likely leads to cell lethality.

Removed in the revised version: We showed above that NAD+ precursors (NAM and NMN) were able to rescue cytotoxicity of PARGi in PARG KO cells, probably by NAD+ depletion and also that these precuresors can also inhibit PARP1/2 activities. Consistently, PARG cKO cells could survive in the presence of NAD+ precursors (NAM and NMN) (Fig. 7E). Taken together, PARG is an essential gene and our PARG KO cells treated with PARGi can mimic the complete loss of PARG activity. These data agree with the genetically engineered mouse models as we discussed above13.

Removed in the revised version: We showed that S phase pADPr signals were specifically detected in PARG KO cells treated with PARGi, which likely originated from unligated Okazaki fragments (Fig. 3). As expectedly, we observed S phase-specific pADPr in PARG cKO cells under double-thymidine block and release condition (Fig. 7F). Moreover, the levels of chromatin-bound PARP1 and pADPr signal, which reflects PARylated PARP1 and/or other PARP1 substrates, increased in PARG cKO cells (Fig. 7G). Treatment with PARPi decreased pADPr signals but increased chromatin-bound PARP1 (Fig. 7G), which is consistent with the results obtained in PARG KO cells (Fig. 4A). Together, PARG, as an essential gene, suppresses initial S phase pADPr signaling and PARP1 chromatin accumulation, which eventually lead to cell death caused by uncontrolled pADPr accumulation throughout cell cycle.

Removed in the revised version: Discussion

Removed in the revised version: Aberrant accumulation of pADPr and/or NAD+ exhaustion disrupts multiple cell processes, including DNA repair, replication stress response, and transcription regulation; therefore, it is highly toxic and eventually leads to cell death10, 22, 42, 43, 62. Although ADP-ribose hydrolases, especially PARG, have a protective effect against excessive PARP1 engagement, it remains unclear when and how PARG activity is engaged during normal cell proliferation. In this study, we showed that the major function of PARG is to regulate S phase-specific pADPr by PARP1, which is activated by unligated Okazaki fragments. When PARG activity is inhibited, these uncontrolled pADPr would eventually result in DNA damage and cell death. Using unbiased FACS- and cell viability-based genome-wide CRISPR screens, we uncovered additional pADPr modulators, which include the genes involved in BER (POLB, XRCC1, and LIG3), Okazaki maturation (LIG1 and FEN1), ARH3, and HPF1. More importantly, we determined that PARG expression is a critical biomarker for PARGi sensitivity. Furthermore, we showed that PARG is an essential gene by suppressing pADPr accumulation. Taken together, our study uncovers a normal function of PARG in proliferating cells: it removes pADPr generated by PARP1 at unligated Okazaki fragments during DNA replication. Of course, this normal function of PARG is also required under conditions of excessive PARP1 engagement because of defects in Okazaki fragment maturation and/or BER.

Removed in the revised version: PARylation is a reversible post-translational modification. In this study, we showed that timely and efficient dePARylation is critical for cell survival. Thus, both PARylation and dePARylation play important functions in the cell. Particularly, dePARylation appears to be indispensable, since it is required for cell survival. The successful clinical application of PARylation inhibition, such as PARPis for the treatment of BRCA-deficient cancers, raises the possibility of targeting dePARylation as an alternative strategy, especially for the treatment of PARPi-resistant tumors. Given that PARG accounts for more than 90% of dePARylation activity7, 29, selective PARG inhibitors are being developed as potential anti-cancer agents. Although defective HR and DDR have been shown to be synthetic lethal with PARG inhibition, comprehensive investigations of cellular response to PARG inhibition have not yet been conducted. In this study, we showed that unexpectedly, PARG is the top synthetic lethality gene in the CRISPR screening with our homemade DDR sgRNA library, which we further validated in HEK293A and HeLa PARG KO cells (Fig. 1B and S1D). Our PARG KO cells plus PARGi mimic complete loss of PARG activity (Fig. S9), which provides is an unique opportunity to study the major function of PARG in proliferating cells. These data were further confirmed with the use of PARG cKO cells we generated (Fig. 7).

Removed in the revised version: PARP1 plays important roles during normal S phase, especially at DNA replication forks and two-ended DSBs47. More recently, unligated Okazaki fragments were identified as the source of SSBs that activate PARP1 in normal proliferating cells6. Moreover, unligated Okazaiki fragments, on which PARP1 may be trapped, can be converted into structures that result in synthetic lethality in HR-defective cancer cells63. PARG is also enriched at DNA replication forks through pADPr- and PCNA-dependent mechanisms14, 64, 65. In addition, PARG may protect DNA replication forks during normal and damage-treated S-phase15, 31, 66. Interestingly, the identification of unligated Okazaki fragments as the source of endougenous pADPr in normal proliferating cells was only noticeable under short-term PARG inhibition. Treatment with another recently developed PARG inhibitor COH34, also led to increased PARylation at replication sites in S phase cells26. These previous reports agree with the data presented in this study. Indeed, we showed robust and specific PARP1/2-dependent S phase pADPr in PARG KO cells treated with PARGi (Fig. 3A and S3). We further confirmed that the S phase-specific pADPr requires normal S phase progression and the generation of Okazaki fragments (Fig. 3B-C and S3B). Moreover, the cytotoxicity of PARGi requires S phase progression (Fig. 3D-E). In addition, PARG KO cells displayed more chromatin-bound PARP1 and PARylated PARP1 (Fig. 4A and S4A). Enhanced pADPr levels were observed in the chromatin fraction of PARG KO cells under PARGi treatment (Fig. 4A). One possible explanation is that in PARG KO cells treated with PARGi, chromatin-bound or trapped PARP1 may not be removed appropriately or efficiently, which further prevents Okazaki fragment ligation/maturation that eventually results in cell lethality. However, the underlying mechanism of more chromatin-bound PARP1 and PARylated PARP1 in PARG KO cells still needs further elucidation.

Removed in the revised version: The inability to remove S phase-specific pADPr eventually led to pADPr throughout cell cycle and γH2AX signaling (Fig. 4C-D and S4D-F), indicating that it is essential to resolve S phase-specific pADPr to maintain cell viability. In agreement with the literature6, 26, we observed S phase-specific pADPr in HeLa cells following PARGi treatment, although such S phase-specific pADPr was not detected in HEK293A cells following PARGi treatment. The lower level of PARG expression may be the reason for this difference between HEK293A and HeLa cells (Fig. S4E). However, significant amounts of PARG activity still existed in HeLa cells, since these cells did not show noticeable PARGi sensitivity unless we knockout or knockdown PARG.

Removed in the revised version: Cell viability–based genome-wide CRISPR/Cas9 screening unbiasedly uncovers sensitivity and resistance genes under certain conditions36, while FACS-based CRISPR/Cas9 screening using reporter cell lines or antibodies recognizing specific signaling molecules can reveal key regulators of signaling pathways37. In this study, we performed both types of screens, which revealed key pADPr regulators, including BER pathway proteins (POLB, XRCC1, LIG3, and CHD1L), DNA replication/Okazaki fragment maturation proteins (RFCs, FEN1, and LIG1), and ARH3 (Fig. 5B and S5A). In addition, consistent with our hypothesis that pADPr accumulation leads to cell death, we showed that loss of genes involved in Okazaki fragment maturation or BER resulted in PARGi sensitivity (Fig. 5C-G). All of these results further validated our working model that S phase-specific pADPr likely originated from unligated Okazaki fragments, which if not removed would eventually lead to DNA damage and cell death.

Removed in the revised version: During DNA replication, Okazaki fragments are normally ligated by LIG1. However, some unligated Okazaki fragments may activate PARP1 and recruit XRCC1/LIG36, which agrees with the results of early studies indicating that both LIG1 and LIG3 are required for Okazaki fragment maturation46. Indeed, we showed that LIG1 depletion, which would result in more DNA nicks in S phase cells, led to the enhancement of S phase PARylation (Fig. S5E). As for LIG3, it can substitute LIG1 function for the ligation of Okazaki fragments46, 67. However, it also has a critical and essential function in mitochondria but is separated for its role in XRCC1-dependent SSBR68. We observed minor S phase pADPr due to LIG3 knockdown with PARGi treatment (Fig. S5E), which is likely due to insufficient KD of LIG3 in these experiments. Unlike LIG1 KO and LIG3 KD cells, XRCC1 KO and POLB KO cells increase PARylation in a cell cycle–independent manner (Fig. S5E), which agrees with the finding that POLB and XRCC1 are required for BER/SSB repair, but not specifically for Okazaki fragment ligation/maturation. Our results suggest that PARylation and dePARylation are mainly involved in two cellular processes, i.e. Okazaki fragment maturation and BER/SSB repair, since both processes have nicked DNA as intermediates, which are likely the physiological substrates that recruit and activate PARP1/2.

Removed in the revised version: Interestingly, two well-known PARylation regulators, ARH3 and HPF1, showed synthetic lethality with PARGi in PARG KO cells (Fig. 5C). Given that ARH3 and PARG are the primary dePARylation enzymes10, it should be predictable that ARH3 was listed as the top hit with PARGi in PARG KO cells in both pADPr signal screen and synthetic lethality screen (Fig. 5B and C). As for HPF1, a previous report showed that the loss of HPF1 will release the activity of PARP1 to PARylation in acidic resides57; moreover, a recent study shows that HPF1 and nucleosomes mediate a dramatic switch in the activity of PARP1 from polymerase to hydrolase69. More important, HPF1 promotes the ligation of the Okazaki fragment by LIG3-XRCC1 as a backup pathway70, while the unligated Okazaki fragments were the source of the S phase pADPr signal in this study.

Removed in the revised version: More importantly, our data indicate that PARG expression is a potential biomarker for PARGi sensitivity. We used the available database and cell lines to validate that cells with low PARG expression are sensitive to PARGi. In addition, we examined PARG expression by the IHC assay in commercial TMA samples and uncovered a fraction of breast and ovarian cancers with no detectable or low PARG expression (Fig. 6). We also examined the sensitivity of HR-deficient cells to PARGi. Although HR deficiency sensitizes cells to PARGi, our results indicate that PARG loss or low expression is still the main driver of PARGi sensitivity in these cells (Fig. S6,7).

Removed in the revised version: Mechanically, we showed that the dePARylation activity of PARG is indispensable for cell survival (Fig. 7 and S9). The residual PARG dePARylation activity observed in PARG KO cells likely supports cell growth, which can be further inhibited by PARGi (Fig. S9). Although a dose-dependent inhibition of PARG activity by PARGi was also noted in wild-type cells, significant fraction of PARG activity remained even in the presence of highest concentration of PARGi used in this study, which is consistent with our results that HEK293A cells are insensitive to PARGi. More importantly, these data further indicate that PARG expression/activity is a potential biomarker for PARGi sensitivity.

Removed in the revised version: Taken together, the dramatic S phase-specific pADPr signaling detected under our experimental conditions provides us a rare opportunity to gain a better understanding of PARG function and pADPr signaling in normal proliferating cells without any exogenous DNA damage. In addition, our systematic CRISPR/Cas9 screens uncover key regulators of pADPr signaling that also contribute to PARGi sensitivity. Moreover, our finding that PARG expression may a potential biomarker of PARGi sensitivity will allow the further development of these inhibitors for the treatment of tumors with PARG loss and therefore offer a new targeted strategy for cancer therapy.

Removed in the revised version: Competing interest statement

Removed in the revised version: The authors declare that they have no conflicts of interest with the content of this article.

Supporting information